Hundreds of myosin 10s are pushed to the tips of filopodia and could cause traffic jams on actin

Author response:

The following is the authors’ response to the previous reviews.

Public Reviews:

Reviewer #1 (Public Review):

Summary:

The manuscript proposes an alternative method by SDS-PAGE calibration of Halo-Myo10 signals to quantify myosin molecules at specific subcellular locations, in this specific case filopodia, in epifluorescence datasets compared to the more laborious and troublesome single molecule approaches. Based on these preliminary estimates, the authors developed further their analysis and discussed different scenarios regarding myosin 10 working models to explain intracellular diffusion and targeting to filopodia.

Strengths:

I confirm my previous assessment. Overall, the paper is elegantly written and the data analysis is appropriately presented. Moreover, the novel experimental approach offers advantages to labs with limited access to high-end microscopy setups (super-resolution and/or EM in particular), and the authors proved its applicability to both fixed and live samples.

Weaknesses:

Myself and the other two reviewers pointed to the same weakness, the use of protein overexpression in U2OS. The authors claim that Myosin10 is not expressed by U2OS, based on Western blot analysis. Does this completely rule out the possibility that what they observed (the polarity of filopodia and the bulge accumulation of Myo10) could be an artefact of overexpression? I am afraid this still remains the main weakness of the paper, despite being properly acknowledged in the Limitations.

Respectfully, our observations do not capture an “artefact” of overexpression but rather the “response” to overexpression. Our goal in this project was to overexpress Myo10 in a situation where it is the limiting reagent for generating filopodia. As Reviewer 3 notes below, overexpression shows that filopodial tips “can accommodate a surprisingly (shockingly) large number of motors.” This is exactly the point. Reviewer 2 considered our handling of this issue to be a strength of the paper. As far as whether bulges occur in endogenous Myo10 systems, please see our comments to Reviewer 3.

I consider all the remaining issues I expressed during the first revision solved.

Reviewer #2 (Public Review):

Summary:

The paper sought to determine the number of myosin 10 molecules per cell and localized to filopodia, where they are known to be involved in formation, transport within, and dynamics of these important actin-based protrusions. The authors used a novel method to determine the number of molecules per cell. First, they expressed HALO tagged Myo10 in U20S cells and generated cell lysates of a certain number of cells and detected Myo10 after SDS-PAGE, with fluorescence and a stained free method. They used a purified HALO tagged standard protein to generate a standard curve which allowed for determining Myo10 concentration in cell lysates and thus an estimate of the number of Myo10 molecules per cell. They also examined the fluorescence intensity in fixed cell images to determine the average fluorescence intensity per Myo10 molecule, which allowed the number of Myo10 molecules per region of the cell to be determined. They found a relatively small fraction of Myo10 (6%) localizes to filopodia. There are hundreds of Myo10 in each filopodia, which suggests some filopodia have more Myo10 than actin binding sites. Thus, there may be crowding of Myo10 at the tips, which could impact transport, the morphology at the tips, and dynamics of the protrusions themselves. Overall, the study forms the basis for a novel technique to estimate the number of molecules per cell and their localization to actin-based structures. The implications are broad also for being able to understand the role of myosins in actin protrusions, which is important for cancer metastasis and wound healing.

Strengths:

The paper addresses an important fundamental biological question about how many molecular motors are localized to a specific cellular compartment and how that may relate to other aspects of the compartment such as the actin cytoskeleton and the membrane. The paper demonstrates a method of estimating the number of myosin molecules per cell using the fluorescently labeled HALO tag and SDS-PAGE analysis. There are several important conclusions from this work in that it estimates the number of Myo10 molecules localized to different regions of the filopodia and the minimum number required for filopodia formation. The authors also establish a correlation between number of Myo10 molecules filopodia localized and the number of filopodia in the cell. There is only a small % of Myo10 that tip localized relative to the total amount in the cell, suggesting Myo10 have to be activated to enter the filopodia compartment. The localization of Myo10 is log-normal, which suggests a clustering of Myo10 is a feature of this motor.

One of the main critiques of the manuscript was that the results were derived from experiments with overexpressed Myo10 and therefore are hard to extrapolate to physiological conditions. The authors counter this critique with the argument that their results provide insight into a system in which Myo10 is a limiting factor for controlling filopodia formation. They demonstrate that U20S cells do not express detectable levels of Myo10 (supplementary Figure 1E) and thus introducing Myo10 expression demonstrates how triggering Myo10 expression impacts filopodia. An example is given how melanoma cells often heavily upregulate Myo10.

In addition, the revised manuscript addresses the concerns about the method to quantitate the number of Myo10 molecules per cell and therefore puncta in the cell. The authors have now made a good faith effort to correct for incomplete labeling of the HALO tag (Figure 2A-C, supplementary Figure 2D-E). The authors also address the concerns about variability in transfection efficiency (Figure 1D-E).

A very interesting addition to the revised manuscript was the quantitation of the number of Myo10 molecules present during an initiation event when a newly formed filopodia just starts to elongate from the plasma membrane. They conclude that 100s of Myo10 molecules are present during an initiation event. They also examined other live cell imaging events in which growth occurs from a stable filopodia tip and correlated with elongation rates.

Weaknesses:

The authors acknowledge that a limitation of the study is that all of the experiments were performed with overexpressed Myo10. They address this limitation in the discussion but also provide important comparisons for how their work relates to physiological conditions, such as melanoma cells that only express large amounts of Myo10 when they are metastatic. Also, the speculation about how fascin can outcompete Myo10 should include a mechanism for how the physiological levels of fascin can complete with the overabundance of Myo10 (page 10, lines 401-408).

We have expanded the discussion about fascin competing with high concentrations of Myo10 in filopodial tips on pg. 15. The key feature is that fascin binding in a bundle is essentially irreversible, so it wins if any space opens up and it manages to bind before the next Myo10 arrives.

Reviewer #3 (Public Review):

Summary

The work represents progress in quantifying the number of Myo10 molecules present in the filopodia tip. It reveals that cells overexpressing fluorescently labeled Myo10 that the tip can accommodate a wide range of Myo10 motors, up to hundreds of molecules per tip.

The revised, expanded manuscript addresses all of this reviewer's original comments. The new data, analysis and writing strengthen the paper. Given the importance of filopodia in many cellular/developmental processes and the pivotal, as yet not fully understood role of Myo10 in their formation and extension, this work provides a new look at the nature of the filopodial tip and its ability to accommodate a large number of Myo10 motor proteins through interactions with the actin core and surrounding membrane.

Specific comments -

(1) One of the comments on the original work was that the analysis here is done using cells ectopically expressing HaloTag-Myo10. The author's response is that cells express a range of Myo10 levels and some metastatic cancer cells, such as breast cancer, have significantly increased levels of Myo10 compared to non-transformed cell lines. It is not really clear how much excess Myo10 is present in those cells compared to what is seen here for ectopic expression in U2OS cells, making a direct correspondence difficult.

We agree, a direct correspondence is difficult, and is further complicated by other variables (e.g., expression levels of Myo10 activators, cargoes, fascin, or other filopodial components) that may differ among cell lines. Properly sorting this out will require additional work in a few key cellular systems.

However, there are two points to keep in mind that somewhat mitigate this concern. First, because ectopic expression of Myo10 causes an ~30x increase in the number of filopodia, the activated Myo10 population is divided over that larger filopodial population. Second, the log-normal distribution of Myo10 across filopodia has a long tail, which means that some cells with low levels of Myo10 will concentrate that Myo10 in a few filopodia.

In response to comments about the bulbous nature of many filopodia tips the authors point out that similar-looking tips are seen when cells are immunostained for Myo10, citing Berg & Cheney (2002). In looking at those images as well as images from papers examining Myo10 immunostaining in metastatic cancer cells (Arjonen et al, 2014, JCI; Summerbell et al, 2020, Sci Adv) the majority of the filopodia tips appear almost uniformly dot-like or circular. There is not too much evidence of the elongated, bulbous filopodial tips seen here.

Yes, the tips in Berg and Cheney are circular, but their size varies considerably (just as a balloon is roughly circular, its size varies with the amount of air it contains). Non-bulbous filopodial tips have a theoretical radius of ~100 nm, which is below the diffraction limit. However, many of the filopodial tips are larger than the diffraction limit in Berg and Cheney, Fig. 1a. We cropped and zoomed in the images to show each fully visible filopodial tip

We attempted to perform a similar analysis of the images in Arjonen and Summerbell. Unfortunately, their images are too small to do so.

However, in reconsidering the approach and results, it is the case that the finding here do establish the plasticity of filopodia tips that can accommodate a surprisingly (shockingly) large number of motors. The authors discuss that their results show that targeting molecules to the filopodia tip is a relatively permissive process (lines 262 - 274). That could be an important property that cells might be able to use to their advantage in certain contexts.

(2) The method for arriving at the intensity of an individual filopodium puncta (starting on line 532 and provided in the Response), and how this is corrected for transfection efficiency and the cell-to-cell variation in expression level is still not clear to this reviewer. The first part of the description makes sense - the authors obtain total molecules/cell based on the estimation on SDS-PAGE using the signal from bound Halo ligand. It then seems that the total fluorescence intensity of each expressing cell analyzed is measured, then summed to get the average intensity/cell. The 'total pool' is then arrived at by multiplying the number of molecules/cell (from SDS-PAGE) by the total number of cells analyzed. After that, then: 'to get the number of molecules within a Myo10 filopodium, the filopodium intensity was divided by the bioreplicate signal intensity and multiplied by 'total pool.' ' The meaning of this may seem simple or straightforward to the authors, but it's a bit confusing to understand what the 'bioreplicate signal intensity' is and then why it would be multiplied by the 'total pool'. This part is rather puzzling at first read.

We agree, such information is critical. We have now revised this description with more precise terms and have included a formula on pg. 20.

Since the approach described here leads the authors to their numerical estimates every effort should be made to have it be readily understood by all readers. A flow chart or diagram might be helpful.

We have added a diagram of the calculations to the supplemental material (Figure 1—figure supplement 3). We hope that both changes will make it easier for others to follow our work.

(3) The distribution of Myo10 punctae around the cell are analyzed (Fig 2E, F) and the authors state that they detect 'periodic stretches of higher Myo10 density along the plasma membrane' (line 123) and also that there is correlation and anti-correlation of molecules and punctae at opposite ends of the cells.

In the first case, it is hard to know what the authors really mean by the phrase 'periodic stretches'. It's not easy to see a periodicity in the distribution of the punctae in the many cells shown in Supp Fig 3. Also, the correlation/anti-correlation is not so easily seen in the quantification shown in Fig 2F. Can the authors provide some support or clarification for what they are stating?

The periodic pattern that we refer to is most apparent in the middle panels of Fig. 2E, F. These panels show the density of Myo10 puncta. These puncta numbers closely correspond to filopodia counts, with the caveat that some filopodia might have multiple puncta. This periodic density might not be as apparent in the raw data shown in Supp. Fig. 3. We have therefore rewritten this paragraph to clarify our observations (pg. 6).

(4) The authors are no doubt aware that a paper from the Tyska lab that employs a completely different method of counting molecules arrives at a much lower number of Myo10 molecules at the filopodial tip than is reported here was just posted (Fitz & Tyska, 2024, bioRxiv, DOI: 10.1101/2024.05.14.593924).

While it is not absolutely necessary for the authors to provide a detailed discussion of this new work given the timing, they may wish to consider adding a note briefly addressing it.

We are aware of this manuscript and that it uses a different approach for calibrating the fluorescence signal in microscopy. However, we are not comfortable commenting on that manuscript at this time, given that it has not yet been peer reviewed with the chance for author revisions.

Recommendations for the authors:

Reviewer #1 (Recommendations For The Authors):

The manuscript the authors are now presenting does not comply with the formatting limits of a Short report, but it is instead presented as a full article type. I believe the authors could shorten the Discussion, and meet the criteria for a more appropriate Short Report format.

For instance, I continue to believe that the study of truncation variants could sustain the claim that membrane binding represents the driving force that leads to Myo10 accumulation. I understand the authors want to address these mechanisms in a follow-up story, for this reason, I encourage them to shorten the discussion, which seems unnecessarily long for a technique-based manuscript.

In the first round of review, Reviewer 3 asked us to expand the discussion. Given that, we are happy with where we have landed on the length of the discussion.

Figure 2, could include some images to facilitate the readers on the different messages of the two rose plots E and F, by picking one of the examples from the supplementary Figure 3

We have now added a supplemental figure showing an example cell (Fig. 2 figure supplement 2). But please note that the averaging of ~150 cells (Fig. 2E, F) should be more reliable to show these overall trends.

Reviewer #2 (Recommendations For The Authors):

Also, the speculation about how fascin can outcompete Myo10 should include a mechanism for how the physiological levels of fascin can complete with the overabundance of Myo10 (page 10, lines 401-408).

As noted above, we have now clarified this point.

Reviewer #3 (Recommendations For The Authors):

line 495 - what is GOC?

We have now defined this oxygen scavenger system in the main text.

lines 603/604 - it is stated that 'velocity analysis does not only account for Myo10 punctum that moved away from the starting point of the trajectory.' It's not clear what this really means.

The sentence now reads: "For Figure 4 parts G-H, note that velocity analysis includes a few Myo10 puncta that switch direction within a single trajectory (e.g., a retracting punctum that then elongates)."

References #4 and #14 are the same.

Thank you for catching that; it has now been corrected.

Fig 1C - the plot for signal intensity versus fmol of protein has numbers for the standard and then live and fixed cells. While the R2 value is quite good, it seems a bit odd that the three (?) data points for live cells are all quite small relative to the fixed cells and all bunched together at the left side of the plot.

As mentioned in the main text, the time post-transfection has a noticeable effect on the level of Myo10 expression. The three fixed-cell bioreplicates had higher Myo10 expression because they were analyzed 48 hours post-transfection compared to the three live-cell bioreplicates (24 hours). Therefore, the fixed cell data points are larger in value because they represent more molecules, and the live cell data points are on the left side of the plot because they represent fewer molecules.

Changes from the earlier version to the revised version

Abstract

Abstract

Added in the revised version: Impact Statement

Added in the revised version: Improved technology for counting proteins captures myosin-10 dynamics throughout the filopodial lifecycle, defines requirements for initiating filopodia, and reveals a crowded filopodial tip where motors frequently exceed available actin.

Earlier version: Myosin 10 (Myo10) is a motor protein well known for its role in filopodia formation. Although Myo10-driven filopodial dynamics have been characterized, there is no information about the absolute number of Myo10 molecules during the filopodial lifecycle. To better understand molecular stoichiometries and packing restraints in filopodia, we measured Myo10 abundance in these structures. Here we combined SDS-PAGE densitometry with epifluorescence microscopy to quantitate HaloTag-labeled Myo10 in U2OS cells. About 6% of total intracellular Myo10 localizes to filopodia, where it is enriched at opposite ends of the cell. Hundreds of Myo10 are found in a typical filopodium, and their distribution across filopodia is log-normal. Some filopodial tips even contain more Myo10 than accessible binding sites on the actin filament bundle. Live-cell movies reveal a dense cluster of over a hundred Myo10 molecules that initiates filopodial elongation. Hundreds of Myo10 molecules continue to accumulate during filopodial growth, but that accumulation ceases when filopodia begin to retract. Rates of filopodial elongation, second-phase elongation, and retraction are inversely related to Myo10 quantities. Our estimates of Myo10 molecules in filopodia provide insight into the physics of packing Myo10, its cargo, and other filopodia-associated proteins in narrow membrane compartments. Our protocol provides a framework for future work analyzing Myo10 abundance and distribution upon perturbation.

Revised version: Myosin 10 (Myo10) is a motor protein known for its role in filopodia formation. Although Myo10-driven filopodial dynamics have been characterized, there is no information about the absolute number of Myo10 molecules during the filopodial lifecycle. To better understand molecular stoichiometries and packing restraints in filopodia, we measured Myo10 abundance in these structures. We combined SDS-PAGE densitometry with epifluorescence microscopy to quantitate HaloTag-labeled Myo10 in U2OS cells. About 6% of total intracellular Myo10 localizes to filopodia, where it enriches at opposite cellular ends. Hundreds of Myo10s are in a typical filopodium, and their distribution across filopodia is log-normal. Some filopodial tips even contain more Myo10 than accessible binding sites on the actin filament bundle. Live-cell movies reveal a dense cluster of over a hundred Myo10 molecules that initiates filopodial elongation. Hundreds of Myo10 molecules continue to accumulate during filopodial growth, but accumulation ceases when retraction begins. Rates of filopodial elongation, second-phase elongation, and retraction are inversely related to Myo10 quantities. Our estimates of Myo10 molecules in filopodia provide insight into the physics of packing Myo10, its cargo, and other filopodia-associated proteins in narrow membrane compartments. Our protocol provides a framework for future work analyzing Myo10 abundance and distribution upon perturbation.

Introduction

Myosins are a group of motor proteins that travel along the cell’s dynamic cytoskeletal highways, binding actin and using ATP as fuel. Apart from myosin 2, the other >40 classes are considered ‘unconventional’ and often include membrane and cargo-binding tail domains1. A myosin of particular interest is myosin 10 (Myo10), a motor protein well known for its role in cellular protrusions2,3. These protrusions, termed filopodia, comprise tightly packed, parallel actin filaments and participate in a multitude of processes such as phagocytosis, directed cell migration, growth-cone guidance, and cell-cell adhesion3. Myo10 has key implications in health dysregulation4. For example, upregulated Myo10 is tied to increased genomic instability5 and breast cancer aggressiveness6.

Earlier version: Myo10’s role in filopodia has been widely investigated2,7–18. Myo10 expression increases dorsal filopodia, and the cargo-binding tail domains, MyTH4 and FERM, are crucial for filopodia formation9. When Myo10 is overexpressed, it produces many Myo10 tip-localized filopodia9,19,20. Estimations on Myo10-positive filopodia length, average number of filopodia per area, and the velocities of extension and retraction have been reported9,2,19,21,22. However, questions regarding the quantity and distribution of Myo10 molecules within the cell still linger. Published images of Myo10 show its prominent localization in the filopodial tip, but a pool of Myo10 in the cell body remains23,24. A quantitative understanding of Myo10 localization could provide further insight into filopodial dynamics, how the crowded environment inside a thin filopodium affects the composition of the filopodial tip complex, and ultimately how Myo10 dysregulation connects to pathology.

Revised version: Myo10’s role in filopodia has been widely investigated2,4,7–17. Myo10 expression increases dorsal filopodia, and the cargo-binding tail domains, MyTH4 and FERM, are crucial for filopodia formation9. When Myo10 is overexpressed, it produces many Myo10 tip-localized filopodia9,18,19. Estimations on Myo10-positive filopodia length, average number of filopodia per area, and the velocities of extension and retraction have been reported9,2,18,20,21. However, questions regarding the quantity and distribution of Myo10 molecules within the cell still linger. Published images of Myo10 show its prominent localization in the filopodial tip, but a pool of Myo10 in the cell body remains22,23. A quantitative understanding of Myo10 localization could provide further insight into filopodial dynamics, how the crowded environment inside a thin filopodium affects the composition of the filopodial tip complex, and ultimately how Myo10 dysregulation connects to pathology.

Earlier version: Studies estimating protein numbers in subcellular compartments have been conducted before25,26–28. Methods include super-resolution microscopy28, mass spectrometry29,30, quantitative western blots25, and photobleaching experiments31. We describe here a simple imaging and analysis method exploiting HaloTag labeling technology applied to Myo10 and filopodia. Similar to quantitative Western blotting of GFP25, our strategy relies on SDS-PAGE densitometry using a fluorescent protein standard to estimate the mean number of Myo10 molecules per cell. In parallel fluorescence microscopy work, we record the fluorescent intensity per cell and convert the per-pixel fluorescent signal to the local number of molecules.

Revised version: Studies estimating protein numbers in subcellular compartments have been conducted before24,25–27. Methods include super-resolution microscopy27, mass spectrometry28,29, quantitative western blots24, and photobleaching experiments30. We describe here a simple imaging and analysis method exploiting HaloTag labeling technology applied to Myo10 and filopodia. Similar to quantitative Western blotting of GFP24, our strategy relies on SDS-PAGE densitometry using a fluorescent protein standard to estimate the mean number of Myo10 molecules per cell. In parallel fluorescence microscopy work, we record the fluorescent intensity per cell and convert the per-pixel fluorescent signal to the local number of molecules.

We find that the bulk of Myo10 remains in the cell body, with hundreds of Myo10s in each filopodium. Myo10 is unevenly distributed across a cell’s filopodia, and some filopodial tips have an excess of Myo10 over accessible actin filament binding-sites. Filopodial initiation from the membrane occurs from a median-sized cluster of 160 Myo10 molecules. Hundreds of Myo10 molecules continue to accumulate in filopodial puncta during filopodial extension phases, but accumulation ceases during filopodial retraction. Nascent filopodial extension, second-phase extension, and retraction rates all diminish with increasing Myo10 amounts, suggesting that a large number of motors suppress filopodial dynamics. Having the number of Myo10 molecules contextualizes interactions of Myo10 with its cargo, the plasma membrane and actin, building a picture of a densely-packed filopodial tip compartment.

Results

HaloTag-based cellular protein quantitation for microscopy

Earlier version: We expressed human HaloTag-Myo10 in U2OS cells to visualize its behavior in filopodia. Wildtype U2OS cells produce few filopodia, but exogenous expression of Myo10 induces abundant surface-attached filopodia (Fig. 1A). This induction suggests that Myo10 is a limiting reagent for filopodia in U2OS cells and allows us to assess the impact of variable Myo10 expression levels. HaloTag labeling was selected due to its high-efficiency binding and the commercial availability of a HaloTag standard protein with a known concentration. Our human, full-length Myo10 construct has an N-terminal HaloTag and C-terminal Flag-tag. Both N-terminal and C-terminal labeled Myo10 have nearly identical fluorescence staining patterns (Supplementary Fig. 1A).

Revised version: We expressed human HaloTag-Myo10 in U2OS cells to visualize its behavior in filopodia. Wildtype U2OS cells produce few filopodia, but exogenous expression of Myo10 induces abundant surface-attached filopodia (Figure 1A). This induction suggests that Myo10 is a limiting reagent for filopodia in U2OS cells and allows us to assess the impact of variable Myo10 expression levels. HaloTag labeling was selected due to its high-efficiency binding and the commercial availability of a HaloTag standard protein with a known concentration. Our human, full-length Myo10 construct has an N-terminal HaloTag and C-terminal Flag-tag. Both N-terminal and C-terminal labeled Myo10 have nearly identical fluorescence staining patterns (Figure 1—figure supplement 1A).

Hundreds of thousands of Myo10 monomer molecules found in Myo10-transfected U2OS cells.

(A) Epifluorescence image of exogenously expressed HaloTag-Myo10 in U2OS cells. Actin is labeled with phalloidin-AF633 (magenta). Myo10 is labeled with HaloTag ligand-TMR (green). Scale bar = 10 µm. (B) The top SDS-PAGE lanes show the indicated quantity (in ng) of HaloTag standard protein. The bottom SDS-PAGE lanes from the same gel show 50,000 cells from 6 separate U2OS transient transfections (except bioreplicate 5, indicated by a red asterisk (*), has 10,000 cells). Bioreplicates 1, 2 and 3 are from live-cell analysis, while bioreplicates 4, 5 and 6 are from fixed-cell analysis. Stain-free shows total protein signal, while TMR shows only TMR-HaloTag-Myo10 signal. Signal was integrated for full-length Myo10 (at ∼250 kDa) and any Myo10 aggregated in the wells at the top of the panels. (C) Standard curve for TMR fluorescence signal of HaloTag standard protein (black dots) compared to signal from HaloTag-Myo10 U2OS cells (red dots = fixed cell experiments, blue dots = live cell experiments). The linear fit is y = 24.32x, where the y-intercept is set to 0. R2 = 0.98. Standard error of slope = 1.42. Grey shading indicates the 95% confidence interval. (D) Distribution of the number of Myo10 molecules per fixed cell (N = 150 cell images; min = 39,000, 95% CI: 33,000-46,000; median = 1,000,000, 95% CI: 870,000-1,200,000; max = 21,000,000, 95% CI: 18,000,000-25,000,000; bins = 30). (E) Distribution of the number of Myo10 molecules per live cell as determined by quantification of the first frame of N = 168 cell movies (min = 79,000, 95% CI: 67,000-92,000; median = 450,000, 95% CI: 370,000-520,000; max = 3,300,000, 95% CI: 2,800,000-3,800,000; bins = 30).

Earlier version: To estimate the number of Myo10 molecules within specific cellular compartments, we combined SDS-PAGE and epifluorescence microscopy. First, we loaded lysate from a known number of Myo10-transfected U2OS cells on an SDS-PAGE gel along with known amounts of HaloTag standard protein, a 61 kDa HaloTag-GST fusion protein (Fig. 1B-C, Supplementary Fig. 1C). Because all samples were incubated with TMR-HaloTag ligand, we use the standard protein’s fluorescence emission in the gel to generate a standard curve and estimate the mean number of Myo10 molecules per cell by densitometry. We propagate the 95% CI from the gel standard curve for our subsequent error estimates. In parallel, we measured TMR-HaloTag Myo10 molecules for the same set of transfected U2OS cells using epifluorescence microscopy. From these measurements, we obtain the total background-subtracted fluorescence intensity for each of the ∼50-cell biological replicates (independent transfections on different days) per fixed and live cell experiments. We used these totals to determine the fluorescence signal per molecule by microscopy. Both microscopy and gel samples were labeled with excess HaloTag ligand (Supplementary Fig. 2A, 2D-E) to ensure maximal labeling of Myo10 molecules with no detectable nonspecific labeling (Supplementary Fig. 1D-E). All Myo10 molecule measurements accounted for our measured 90% TMR-HaloLigand labeling efficiency (Supplementary Fig. 2B). Importantly, wildtype U2OS cells do not express detectable levels of Myo10 (Supplementary Fig. 1E) and generate few filopodia. Therefore, essentially all Myo10 molecules carry the HaloTag label, and these molecules are responsible for inducing the filopodia that we observe. We pooled measurements from the three bioreplicates, resulting in the analysis of 150 cells of varying Myo10 expression levels in fixed cell experiments, and 168 cells in live-cell experiments (Fig. 1D-E).

Revised version: To estimate the number of Myo10 molecules within specific cellular compartments, we combined SDS-PAGE and epifluorescence microscopy. First, we loaded lysate from a known number of Myo10-transfected U2OS cells on an SDS-PAGE gel along with known amounts of HaloTag standard protein, a 61 kDa HaloTag-GST fusion protein (Figure 1B-C, Figure 1—figure supplement 1C). Because all samples were incubated with TMR-HaloTag ligand, we use the standard protein’s fluorescence emission in the gel to generate a standard curve and estimate the mean number of Myo10 molecules per cell by densitometry. We propagate the 95% CI from the gel standard curve for our subsequent error estimates. In parallel, we measured TMR-HaloTag Myo10 molecules for the same set of transfected U2OS cells using epifluorescence microscopy. From these measurements, we obtain the total background-subtracted fluorescence intensity for each of the ∼50-cell biological replicates (independent transfections on different days) per fixed and live cell experiments. We used these totals to determine the fluorescence signal per molecule by microscopy. Both microscopy and gel samples were labeled with excess HaloTag ligand (Figure 1—figure supplement 2A, 2D-E) to ensure maximal labeling of Myo10 molecules with no detectable nonspecific labeling (Figure 1—figure supplement 1D-E). All Myo10 molecule measurements accounted for our measured 90% TMR-HaloLigand labeling efficiency (Figure 1—figure supplement 2B). Importantly, wildtype U2OS cells do not express detectable levels of Myo10 (Figure 1—figure supplement 1E) and generate few filopodia. Therefore, essentially all Myo10 molecules carry the HaloTag label, and these molecules are responsible for inducing the filopodia that we observe. We pooled measurements from the three bioreplicates, resulting in the analysis of 150 cells of varying Myo10 expression levels in fixed cell experiments, and 168 cells in live-cell experiments (Figure 1D-E).

Limited quantities of Myo10 in filopodia

Earlier version: We start with observations in fixed cells that report the distribution of Myo10 throughout U2OS cells. Exogenous expression of Myo10 results in 12–116 Myo10-positive filopodia per cell (median: 55; Fig. 2A, Supplementary Fig. 1B). Varying filopodia density likely reflects cells at different stages of migration or signaling2,32. Despite cells containing a median of 1,000,000 (95% CI: 870,000-1,200,000) total Myo10 molecules (Fig. 1D), only a small proportion of Myo10 localizes to filopodia (median: 5.4%; Fig. 2B). This small proportion of filopodial Myo10 is likely due to a limited Myo10 activation signal that is necessary to relieve autoinhibition and begin processive motility along the filopodial shaft.

Revised version: We start with observations in fixed cells that report the distribution of Myo10 throughout U2OS cells. Exogenous expression of Myo10 results in 12–116 Myo10-positive filopodia per cell (median: 55; Figure 2A, Figure 1—figure supplement 1B). Varying filopodia density likely reflects cells at different stages of migration or signaling2,31. Despite cells containing a median of 1,000,000 (95% CI: 870,000-1,200,000) total Myo10 molecules (Figure 1D), only a small proportion of Myo10 localizes to filopodia (median: 5.4%; Figure 2B). This small proportion of filopodial Myo10 is likely due to a limited Myo10 activation signal that is necessary to relieve autoinhibition and begin processive motility along the filopodial shaft.

Only a small portion of intracellular Myo10 is activated and enters filopodia, and Myo10 is unevenly distributed around the cell.

The following values are from fixed cell images. (A) Distribution of the number of Myo10-positive filopodia per cell (N = 8,733 Myo10-positive filopodia, 150 cells, min = 12, median = 55, max = 116). (B) Distribution of the percent of Myo10 localized in the filopodia per cell (N = 150 cells, min = 0.86%, median = 5.35%, max = 19.87%). (C) Correlation between number of Myo10-positive filopodia in a cell and the filopodial number of Myo10 molecules in the cell. The slope of power law function is 0.36. (D) Correlation between number of Myo10-positive filopodia in a cell and the total number of Myo10 molecules in the cell. The slope of power law function is 0.29. (E) Spatial correlation of Myo10-rich regions of the cell edge. Each cell was divided into 20 angular sections, and the section with the most molecules was aligned to 0°. Section quantities were then averaged across cells. Molecules, puncta, and molecules per punctum are shown. (F) Spatial correlation of Myo10-poor regions of the cell edge. As in E, but the section with the fewest molecules was aligned to 0° for each cell’s rose plot. If >1 section contained no Myo10, a randomly selected empty Myo10 section was aligned to 0°. Error bars in E, F are the standard error of the mean for 500 bootstrapped samples of the 150 cells. Note the correlation of both molecules and puncta at opposite ends of cells (0°, 180°), and the anticorrelation with the two sides (90°, 270°).

Earlier version: What does a cell’s total Myo10 filopodial signal indicate about filopodia formation? Higher Myo10 filopodial signal correlates with more filopodia, and ∼100x more Myo10 appears to boost filopodia by 3x (Fig. 2C). Likewise, higher total intracellular Myo10 correlates with more filopodia (Fig. 2D), which supports our earlier hypothesis that Myo10 limits filopodial production in U2OS cells.

Revised version: What does a cell’s total Myo10 filopodial signal indicate about filopodia formation? Higher Myo10 filopodial signal correlates with more filopodia, and ∼100x more Myo10 appears to boost filopodia by 3x (Figure 2C). Likewise, higher total intracellular Myo10 correlates with more filopodia (Figure 2D), which supports our earlier hypothesis that Myo10 limits filopodial production in U2OS cells.

Spatial patterning of Myo10-decorated filopodia

Removed in the revised version: Interactions with membrane-bound proteins, small cytosolic factors, and cortical actin networks might impact where Myo10 is localized. If these influences operate beyond the average separation between filopodia, we might see correlated spatial patterns where Myo10 is concentrated in sets of nearby filopodia and depleted from other sets. Indeed, Myo10 is not uniformly distributed, instead concentrating in cellular zones at opposite ends of the cell (Fig. 2E & F, Supplementary Fig. 3). Cells display periodic stretches of higher Myo10 density along the cell membrane, and some cells even contain membrane regions devoid of filopodia and Myo10 (Supplementary Fig. 3).

Earlier version: Filopodia containing high Myo10 signal are surrounded by more neighboring filopodia, as evidenced by increased numbers of proximal Myo10 puncta (Fig. 2E, center). We define puncta as any cluster/spot of Myo10 detected by segmentation (see Methods for details on image segmentation analysis). The number of Myo10 per filopodium appears random and relatively fixed because the number of molecules per punctum is consistent (Fig. 2E, right). Potentially a local activation signal initiates filopodia at a particular site, and then the signal spreads from the high Myo10 zone to generate more puncta in the immediate vicinity. The opposite is weakly true, by which filopodia with low Myo10 levels tend to be next to fewer Myo10 puncta (Fig. 2F, center). We noticed that a filopodial punctum can contain as few as ten Myo10 molecules (Supplementary Fig. 4A); these dim Myo10 puncta were primarily tip-localized (Supplementary Fig. 4B,C).

Revised version: Interactions with membrane-bound proteins, small cytosolic factors, and cortical actin networks might impact where Myo10 is localized. If these influences operate beyond the average separation between filopodia, we might see correlated spatial patterns of higher and lower Myo10 density along the edge of the cell. Indeed, Myo10 is not uniformly distributed, instead concentrating in cellular zones at opposing sides of the cell (Figure 2E & F, Figure 2—figure supplement 1). Sectors of the cell with the highest quantity of Myo10 have neighboring sectors that are also high (Figure 2E, left). Moreover, these high-signal sectors have more Myo10 puncta, as does the opposite side of the cell. Thus, there is a periodic high-low-high-low pattern of punctum density on the cell’s perimeter (Figure 2E, center; see Figure 2—figure supplement 2 for an example cell). We define puncta as any cluster/spot of Myo10 detected by segmentation (see Methods for details on image segmentation analysis). However, the number of Myo10 per filopodium is relatively constant on average around the cell (Figure 2E, right). Potentially a local activation signal initiates filopodia at a particular site, and then the signal spreads from the high Myo10 zone to generate more puncta in the immediate vicinity. Likewise, sectors of the cell with low Myo10 levels tend to be surrounded by fewer Myo10 puncta (Figure 2F, center). Some cells even contain membrane regions devoid of filopodia and Myo10 (Figure 2—figure supplement 1). We noticed that a filopodial punctum can contain as few as ten Myo10 molecules (Figure 3— figure supplement 1A); these dim Myo10 puncta were primarily tip-localized (Figure 3—figure supplement 1B,C).

Myo10 often saturates accessible actin at filopodial tips

Earlier version: There is a vast range of Myo10 molecules per filopodium (median: 730 molecules, 95% CI: 610-850; Fig. 3A, B), and its distribution among filopodia is apparently log-normal (Supplementary Fig. 4D, E). Because log-normal distributions have long tails, we wondered how filopodia can support the movement of extreme Myo10 levels at filopodial tips. To address this question, we compared Myo10 concentrations to the amount of available actin sites in filopodia. We randomly selected a set of 90 filopodial tip-localized Myo10 puncta and measured the apparent length of the punctum. These lengths vary from ∼250 nm (diffraction limited) to just over 1.5 µm in the case of the most elongated punctum. To estimate Myo10 concentrations, we used published values to define the geometry of a “typical” filopodium 33,34,35. Considering a filopodium to be a cylindrical tube, we calculated local concentrations of Myo10 at filopodial tips using the number of molecules, the length of a Myo10 punctum, and a fixed radius of 100 nm36. At the tips, Myo10 ranges from ∼ 6 μM to 560 μM (Fig. 3C, D).

Revised version: There is a vast range of Myo10 molecules per filopodium (median: 730 molecules, 95% CI: 610-850; Figure 3A, B), and its distribution among filopodia is apparently log-normal (Figure 3—figure supplement 1D, E). Because log-normal distributions have long tails, we wondered how filopodia can support the movement of extreme Myo10 levels at filopodial tips. To address this question, we compared Myo10 concentrations to the amount of available actin sites in filopodia. We randomly selected a set of 90 filopodial tip-localized Myo10 puncta and measured the apparent length of the punctum. These lengths vary from ∼250 nm (diffraction limited) to just over 1.5 µm in the case of the most elongated punctum. To estimate Myo10 concentrations, we used published values to define the geometry of a “typical” filopodium 32,33,34. Considering a filopodium to be a cylindrical tube, we calculated local concentrations of Myo10 at filopodial tips using the number of molecules, the length of a Myo10 punctum, and a fixed radius of 100 nm35. At the tips, Myo10 ranges from ∼ 6 μM to 560 μM (Figure 3C, D).

Hundreds of Myo10 molecules are found in a filopodium, potentially in excess over available actin at the filopodial tip.

Earlier version: The following values are from fixed cell images. (A) Distribution of the number of Myo10 molecules per filopodium (N = 150 cells, 8,733 filopodia; min = 6, 95% CI: 5-7; median = 730, 95% CI: 610-850; max = 80,000, 95% CI: 67,000-93,000; 62 values >20,000 not shown; bins = 100). (B) Cumulative distribution function plot of data in part A. (C) Distribution of the number of Myo10 molecules at the filopodial tip (90 randomly chosen filopodia tip-localized Myo10 puncta from 9 different cell images; min = 66, 95% CI: 55-76; median = 788, 95% CI: 660-915; max =11,000, 95% CI: 9,600-13,000; bins = 30). (D) The local concentration of Myo10 at a filopodial tip. To estimate the volume, we measured the length of the filopodia tip-localized Myo10 puncta from part C in ImageJ. We then modeled filopodium as a cylinder of radius = 0.1 µm (published average). Min = 6.2 µM, 95% CI: 5.2-7.2; median = 84 µM, 95% CI: 70-97; max = 560 µM, 95% CI: 470-650, bins = 30. Blue dashed vertical line indicates the concentration of F-actin accessible for Myo10 binding in a filopodium (∼96 µM). (E) Scatterplots of molecules vs. length for the puncta from part C. The phase boundary shows the 96 µM threshold from part D. (F) As in part E, but the line represents an estimate of allowable Myo10 on the filopodial tip membrane area. See supplementary Fig. 5 for membrane occupancy estimates. (G) Model of a Myo10 traffic jam at the filopodium tip. Not enough available actin monomers results in a population of free Myo10 (in blue). The free Myo10 is detached from actin but potentially still membrane-associated. (H) Model of frayed actin filaments at the filopodium tip. If actin filaments are not neatly packed into parallel bundles at the filopodium tip, disorganized and frayed actin filaments yield more accessible binding sites to Myo10.

Revised version: The following values are from fixed cell images. (A) Distribution of the number of Myo10 molecules per filopodium (N = 150 cells, 8,733 filopodia; min = 6, 95% CI: 5-7; median = 730, 95% CI: 610-850; max = 80,000, 95% CI: 67,000-93,000; 62 values >20,000 not shown; bins = 100). (B) Cumulative distribution function plot of data in part A. (C) Distribution of the number of Myo10 molecules at the filopodial tip (90 randomly chosen filopodia tip-localized Myo10 puncta from 9 different cell images; min = 66, 95% CI: 55-76; median = 788, 95% CI: 660-915; max =11,000, 95% CI: 9,600-13,000; bins = 30). (D) The local concentration of Myo10 at a filopodial tip. To estimate the volume, we measured the length of the filopodia tip-localized Myo10 puncta from part C in ImageJ. We then modeled filopodium as a cylinder of radius = 0.1 µm (published average). Min = 6.2 µM, 95% CI: 5.2-7.2; median = 84 µM, 95% CI: 70-97; max = 560 µM, 95% CI: 470-650, bins = 30. Blue dashed vertical line indicates the concentration of F-actin accessible for Myo10 binding in a filopodium (∼96 µM). (E) Scatterplots of molecules vs. length for the puncta from part C. The phase boundary shows the 96 µM threshold from part D. (F) As in part E, but the line represents an estimate of allowable Myo10 on the filopodial tip membrane area. See Figure 3—figure supplement 2 for membrane occupancy estimates. (G) Model of a Myo10 traffic jam at the filopodium tip. Not enough available actin monomers results in a population of free Myo10 (in blue). The free Myo10 is detached from actin but potentially still membrane-associated. (H) Model of frayed actin filaments at the filopodium tip. If actin filaments are not neatly packed into parallel bundles at the filopodium tip, disorganized and frayed actin filaments yield more accessible binding sites to Myo10.

Earlier version: Interestingly, Myo10 continues to flow into filopodia, even when there may be insufficient actin at the filopodial tip. To estimate the amount of F-actin available for binding, we modeled a filopodium of radius 100 nm comprising 30 actin filaments, of which 16 filaments are on the exposed surface of the bundle. Nagy et al. posited that only 4 of the 13 actin monomers per helical turn are available to Myo10 binding due to steric constraints within a fascin-actin bundle34. Using Zhuralev et al.’s equation35 for calculating F-actin monomer concentration in a filopodium, we estimate that ∼96 µM of actin monomers are available for Myo10 binding (see Supplementary Fig. 5 for full calculations).

Revised version: Interestingly, Myo10 continues to flow into filopodia, even when there may be insufficient actin at the filopodial tip. To estimate the amount of F-actin available for binding, we modeled a filopodium of radius 100 nm comprising 30 actin filaments, of which 16 filaments are on the exposed surface of the bundle. Nagy et al. posited that only 4 of the 13 actin monomers per helical turn are available to Myo10 binding due to steric constraints within a fascin-actin bundle33. Using Zhuralev et al.’s equation34 for calculating F-actin monomer concentration in a filopodium, we estimate that ∼96 µM of actin monomers are available for Myo10 binding (see Figure 3—figure supplement 2 for full calculations).

Earlier version: If a filopodium contains ∼96 μM of F-actin available to Myo10, and up to 560 μM Myo10, then sometimes Myo10 appears in excess (Fig. 3D, E). While excess Myo10 is not attached to actin, it may still be docked on the plasma membrane. In this scenario, we estimate available membrane area to include the curved hemisphere area and additional membrane space occupied by the length of the Myo10 tip-localized puncta in Fig. 3C. We likewise estimate the footprint of the Myo10 C-terminal cargo binding domains, including those involved in membrane binding (see Supplementary Fig. 6 for details). We find that with even relatively large footprint estimates, there is still sufficient membrane area for filopodial tip-localized Myo10 (Fig. 3F). In the four cases where the available membrane is exceeded in Fig. 3F, it is likely because excess Myo10 causes a bulbous extension of the filopodial tip. This extension would provide more area for binding compared to a stereotypical, cylindrical filopodium, since a sphere enclosing a cylinder offers more surface area than a hemispherical cap.

Revised version: If a filopodium contains ∼96 μM of F-actin available to Myo10, and up to 560 μM Myo10, then sometimes Myo10 appears in excess (Figure 3D, E). While excess Myo10 is not attached to actin, it may still be docked on the plasma membrane. In this scenario, we estimate available membrane area to include the curved hemisphere area and additional membrane space occupied by the length of the Myo10 tip-localized puncta in Figure 3C. We likewise estimate the footprint of the Myo10 C-terminal cargo binding domains, including those involved in membrane binding (see Figure 3—figure supplement 3 for details). We find that with even relatively large footprint estimates, there is still sufficient membrane area for filopodial tip-localized Myo10 (Figure 3F). In the four cases where the available membrane is exceeded in Figure 3F, it is likely because excess Myo10 causes a bulbous extension of the filopodial tip. This extension would provide more area for binding compared to a stereotypical, cylindrical filopodium, since a sphere enclosing a cylinder offers more surface area than a hemispherical cap.

Earlier version: This scenario of excess Myo10 at the filopodial tip would inevitably lead to a molecular traffic jam within the filopodia (Fig. 3G, Supplementary Fig. 5), during which Myo10 likely remains bound to the plasma membrane (Fig. 3F, Supplementary Fig. 6). Alternatively, the ends of actin bundles could be frayed at filopodial tips, increasing the available actin binding sites for Myo10 (Fig. 3H).

Revised version: This scenario of excess Myo10 at the filopodial tip would inevitably lead to a molecular traffic jam within the filopodia (Figure 3G, Figure 3—figure supplement 2), during which Myo10 likely remains bound to the plasma membrane (Figure 3F, Figure 3—figure supplement 3). Alternatively, the ends of actin bundles could be frayed at filopodial tips, increasing the available actin binding sites for Myo10 (Figure 3H).

Myo10 accumulation over the filopodial lifecycle

Earlier version: All the above observations are based on fixed U2OS cells and therefore provide a static view of filopodial systems. To understand Myo10 levels and how they change during the dynamic filopodial lifecycle, we adapted our fluorescence calibration method for live cell imaging. We found it critical to rapidly collect 50–60 snapshots of individual cells for the calibration before collecting our movies on each sample. Adequate temperature control of the sample was also crucial, as the U2OS cells would rapidly contract their filopodia in response to cold. We identified and focused on three events in the filopodial lifecycle: filopodial initiation from a Myo10 punctum that appears at the plasma membrane, second-phase elongation from a Myo10 tip punctum in an established filopodium19, and filopodial retraction.

Revised version: All the above observations are based on fixed U2OS cells and therefore provide a static view of filopodial systems. To understand Myo10 levels and how they change during the dynamic filopodial lifecycle, we adapted our fluorescence calibration method for live cell imaging. We found it critical to rapidly collect 50–60 snapshots of individual cells for the calibration before collecting our movies on each sample. Adequate temperature control of the sample was also crucial, as the U2OS cells would rapidly contract their filopodia in response to cold. We identified and focused on three events in the filopodial lifecycle: filopodial initiation from a Myo10 punctum that appears at the plasma membrane, second-phase elongation from a Myo10 tip punctum in an established filopodium18, and filopodial retraction.

Earlier version: Because Myo10 seems to be limiting filopodial production in U2OS cells, we wondered how many Myo10 molecules are needed for filopodial initiation. As Myo10 coalesces at a site on the plasma membrane and prepares to shoot into a filopodium, a median of 160 molecules gather in the punctum (Fig. 4A, Movie 1). The minimal value at the initiation site is 52 molecules, which may represent the lower limit of Myo10 needed for filopodial initiation. Although we attempted to observe Myo10 coalescence at the site of nascent filopodium, we were unable to detect them due to the limited sensitivity of the epifluorescence microscope. We did not use our single-molecule TIRF microscopy because it has a limited field of view that is unable to capture an entire U2OS cell, which is needed to perform the fluorescence calibration.

Revised version: Because Myo10 seems to be limiting filopodial production in U2OS cells, we wondered how many Myo10 molecules are needed for filopodial initiation. As Myo10 coalesces at a site on the plasma membrane and prepares to shoot into a filopodium, a median of 160 molecules gather in the punctum (Figure 4A, Figure 4—video 1). The minimal value at the initiation site is 52 molecules, which may represent the lower limit of Myo10 needed for filopodial initiation.

Myo10 dynamics from live cell movies.

Earlier version: (A-C) Dense Myo10 puncta appear at the start of each phase of the filopodial lifecycle. Distributions of the number of Myo10 molecules in puncta upon: (A) Filopodial initiation (min = 52, 95% CI: 44-61; median = 160, 95% CI: 140-190; max = 1200, 95% CI: 970-1,300). (B) Second-phase elongation (min = 65, 95% CI: 54-75; median = 290, 95% CI: 240-340; max = 1,400, 95% CI: 1,200-1,600). (C) Filopodial retraction (min = 71, 95% CI: 60-82; median = 240, 95% CI: 200-280; max = 1,200, 95% CI: 1,000-1,400). Values are the means of the first two frames after spot detection and identification of filopodial lifecycle stage. Histograms A-C have 30 bins each. (D-F) Accumulation of Myo10 in puncta after filopodial initiation or second-phase elongation, but not after retraction. Evolution of number of molecules for each filopodial phase over time for: (D) Filopodial initiation. (E) Second-phase elongation. (F) Filopodial retraction. Starting values from (A-C) were subtracted from all traces to obtain delta over time. The generalized additive model (GAM) trend lines (blue) exclude long times (> 100 s) with few surviving trajectories. The outlier trajectory indicated by the magenta asterisk is from Movie 3, and the cyan asterisk is from Movie 4. (G-I) Myo10 punctum speeds are inversely correlated with the number of Myo10 molecules. Plots of instantaneous speeds vs. molecules for: (G) Filopodial initiation (min = 0.7, median = 160, max = 2,200 nm/s, Spearman’s ρ = -0.51, p < 2.2*10-16). (H) Second-phase elongation (min speed = 0.4, median = 110, max = 3,000 nm/s, Spearman’s ρ = -0.55, p < 2.2*10-16). (I) Filopodial retraction (min speed = 1.2, median = 140, max = 1,900 nm/s, Spearman’s ρ = -0.45, p < 2.2*10-16). Color signifies Myo10 puncta belonging to the same trajectory within each event type. Colors are independent in each panel. Blue lines are GAM trend lines. Panels A, D, G: 237 trajectories from 31 cells; B, E, H: 51 trajectories from 19 cells; C, F, I: 58 trajectories from 26 cells.

Revised version: (A-C) Dense Myo10 puncta appear at the start of each phase of the filopodial lifecycle. Distributions of the number of Myo10 molecules in puncta upon: (A) Filopodial initiation (min = 52, 95% CI: 44-61; median = 160, 95% CI: 140-190; max = 1200, 95% CI: 970-1,300). (B) Second-phase elongation (min = 65, 95% CI: 54-75; median = 290, 95% CI: 240-340; max = 1,400, 95% CI: 1,200-1,600). (C) Filopodial retraction (min = 71, 95% CI: 60-82; median = 240, 95% CI: 200-280; max = 1,200, 95% CI: 1,000-1,400). Values are the means of the first two frames after spot detection and identification of filopodial lifecycle stage. Histograms A-C have 30 bins each. (D-F) Accumulation of Myo10 in puncta after filopodial initiation or second-phase elongation, but not after retraction. Evolution of number of molecules for each filopodial phase over time for: (D) Filopodial initiation. (E) Second-phase elongation. (F) Filopodial retraction. Starting values from (A-C) were subtracted from all traces to obtain delta over time. The generalized additive model (GAM) trend lines (blue) exclude long times (> 100 s) with few surviving trajectories. The outlier trajectory indicated by the magenta asterisk is from Figure 4—video 3, and the cyan asterisk is from Figure 4—video 4. (G-I) Myo10 punctum speeds are inversely correlated with the number of Myo10 molecules. Plots of instantaneous speeds vs. molecules for: (G) Filopodial initiation (min = 0.7, median = 160, max = 2,200 nm/s, Spearman’s p = −0.51, p < 2.2*10-16). (H) Second-phase elongation (min speed = 0.4, median = 110, max = 3,000 nm/s, Spearman’s p = −0.55, p < 2.2*10-16). (I) Filopodial retraction (min speed = 1.2, median = 140, max = 1,900 nm/s, Spearman’s p = −0.45, p < 2.2*10-16). Color signifies Myo10 puncta belonging to the same trajectory within each event type. Colors are independent in each panel. Blue lines are GAM trend lines. Panels A, D, G: 237 trajectories from 31 cells; B, E, H: 51 trajectories from 19 cells; C, F, I: 58 trajectories from 26 cells.

Added in the revised version: Although we attempted to observe Myo10 coalescence at the site of nascent filopodium, we were unable to detect them due to the limited sensitivity of the epifluorescence microscope. We did not use our single-molecule TIRF microscopy because it has a limited field of view that is unable to capture an entire U2OS cell, which is needed to perform the fluorescence calibration.

Earlier version: Oftentimes, a new filopodial tip can emerge from an existing one in a process that we call second-phase elongation. Such a Myo10-induced multi-cycle filopodial elongation mechanism has been previously proposed15. We find a median of 290 Myo10 molecules organize and form the separate punctum at the start of second-phase elongation (Fig. 4B, Movie 2, Supplementary Fig. 7). This number represents about double the number of Myo10 involved in nascent elongation (median: 160). Retracting Myo10 puncta initially contain a median of 240 molecules (Fig. 4C).

Revised version: Oftentimes, a new filopodial tip can emerge from an existing one in a process that we call second-phase elongation. Such a Myo10-induced multi-cycle filopodial elongation mechanism has been previously proposed14. We find a median of 290 Myo10 molecules organize and form the separate punctum at the start of second-phase elongation (Figure 4B, Figure 4—video 2, Figure 4—figure supplement 1). This number represents about double the number of Myo10 involved in nascent elongation (median: 160). Retracting Myo10 puncta initially contain a median of 240 molecules (Figure 4C).

Earlier version: Both the initial punctum and the second-phase punctum accumulate additional Myo10 molecules over time (Fig. 4D, E). To estimate the accumulated amount, we took the last point of each trajectory and performed a t-test against the null hypothesis of a zero mean. Initial puncta accumulate a mean of 95 molecules (95% CI: 74, 120; t = 8.8; df = 236; p-value = 2.7*10-16), while second-phase puncta accumulate 307 molecules (95% CI: 130, 480; t = 3.5; df = 50; p-value = 0.001). However, unlike these two filopodial growth processes, retracting Myo10 puncta remain relatively constant as they return to the cell body (mean: -73 molecules; 95% CI: -140, - 10; t = -2.3; df = 57; p-value = 0.025; Fig 4F). When retracting filopodia reach the cell body, the punctum often dissipates. However in a few cases, a new Myo10 punctum will coalesce and emerge from the location of a previous one (Movie 3).

Revised version: Both the initial punctum and the second-phase punctum accumulate additional Myo10 molecules over time (Figure 4D, E). To estimate the accumulated amount, we took the last point of each trajectory and performed a t-test against the null hypothesis of a zero mean. Initial puncta accumulate a mean of 95 molecules (95% CI: 74, 120; t = 8.8; df = 236; p-value = 2.7*10-16), while second-phase puncta accumulate 307 molecules (95% CI: 130, 480; t = 3.5; df = 50; p-value = 0.001). However, unlike these two filopodial growth processes, retracting Myo10 puncta remain relatively constant as they return to the cell body (mean: −73 molecules; 95% CI: −140, - 10; t = −2.3; df = 57; p-value = 0.025; Fig 4F). When retracting filopodia reach the cell body, the punctum often dissipates. However in a few cases, a new Myo10 punctum will coalesce and emerge from the location of a previous one (Figure 4—video 3).

Increasing quantities of Myo10 slow the punctum movements in all phases of the filopodial lifecycle

Earlier version: Finally, we wondered how the speeds of filopodial elongation and retraction are related to the number of Myo10 molecules in a tip punctum. Such values would help to constrain models that depend on Myo10 flux to deliver components to the filopodial tip or to generate force at the tip to assist actin polymerization. Interestingly, speeds of tip puncta during filopodial initiation, second-phase elongation, and retraction are all inversely related to the amount of Myo10 in the punctum (Fig. 4G-I). In each case, fewer molecules are associated with a higher variance in speed. The median speed of elongating Myo10 puncta is 160 nm/s, which is similar to previously published values37,38,2,19. Consistent with the trend of more molecules resulting in slower puncta velocities, the larger puncta involved in second-phase elongation display a slower median speed of 110 nm/s. The median speed of retracting Myo10 puncta is 110 nm/sec, similar to previously reported values19. For retracting puncta moving faster than actin retrograde flow, the velocity of Myo10 likely reflects a collapse of the whole filopodium or other complex filopodial dynamics. Our live movies capture some of these non-constant Myo10 retrograde events. Movie 4 exemplifies a decreasing intensity trajectory, wherein a Myo10 punctum rapidly ‘snaps’ closer to the cell body before diffusing into the cytoplasm. The ‘snap’ could be explained by the filopodium undergoing coiling or buckling, possibly induced in part by Myo10-generated force39,40.

Revised version: Finally, we wondered how the speeds of filopodial elongation and retraction are related to the number of Myo10 molecules in a tip punctum. Such values would help to constrain models that depend on Myo10 flux to deliver components to the filopodial tip or to generate force at the tip to assist actin polymerization. Interestingly, speeds of tip puncta during filopodial initiation, second-phase elongation, and retraction are all inversely related to the amount of Myo10 in the punctum (Figure 4G-I). In each case, fewer molecules are associated with a higher variance in speed. The median speed of elongating Myo10 puncta is 160 nm/s, which is similar to previously published values36,37,2,18. Consistent with the trend of more molecules resulting in slower puncta velocities, the larger puncta involved in second-phase elongation display a slower median speed of 110 nm/s. The median speed of retracting Myo10 puncta is 110 nm/sec, similar to previously reported values18. For retracting puncta moving faster than actin retrograde flow, the velocity of Myo10 likely reflects a collapse of the whole filopodium or other complex filopodial dynamics. Our live movies capture some of these non-constant Myo10 retrograde events. Figure 4—video 4 exemplifies a decreasing intensity trajectory, wherein a Myo10 punctum rapidly ‘snaps’ closer to the cell body before diffusing into the cytoplasm. The ‘snap’ could be explained by the filopodium undergoing coiling or buckling, possibly induced in part by Myo10-generated force38,39.

Discussion

Limited activation of Myo10

Earlier version: Myo10 is the limiting reagent for constructing filopodia in U2OS cells, but less than 20% arrives in filopodia when exogenously expressed (Fig. 2B). There is ample time for Myo10 to diffuse from the cytosolic pool to the numerous filopodial sites present in these cells, where it would then be captured. Thus, we expect that most of the cytosolic pool of Myo10 is in an inactive, autoinhibited state. Available PtdIns(3,4,5)P3 (PIP3) in the cell could be limiting the portion of free intracellular Myo10 that is activated and entering filopodia. Umeki et al. proposed that Myo10 binds PIP3 at cell peripheries and dimerizes, becomes activated and converges local actin filaments to form the filopodial base41. Although Myo10 can participate in filopodia formation independently of VASP proteins and substrate attachment9, other proteins have been found to be conserved in mechanisms of filopodia initiation. Components such as Cdc4242, formin43, Arp2/3 complex44, neurofascin32 or VASP45,46 could also be prospective limiting partners of Myo10 entry into filopodia.

Revised version: Myo10 is the limiting reagent for constructing filopodia in U2OS cells, but less than 20% arrives in filopodia when exogenously expressed (Figure 2B). There is ample time for Myo10 to diffuse from the cytosolic pool to the numerous filopodial sites present in these cells, where it would then be captured. Thus, we expect that most of the cytosolic pool of Myo10 is in an inactive, autoinhibited state. Available PtdIns(3,4,5)P3 (PIP3) in the cell could be limiting the portion of free intracellular Myo10 that is activated and entering filopodia. Umeki et al. proposed that Myo10 binds PIP3 at cell peripheries and dimerizes, becomes activated and converges local actin filaments to form the filopodial base40. Although Myo10 can participate in filopodia formation independently of VASP proteins and substrate attachment9, other proteins have been found to be conserved in mechanisms of filopodia initiation. Components such as Cdc4241, formin42, Arp2/3 complex43, neurofascin31 or VASP44,45 could also be prospective limiting partners of Myo10 entry into filopodia.

Earlier version: Prior information on the absolute number of Myo10 molecules in filopodia is sparse. Kerber et al. transfected HeLa cells with low levels of GFP-Myo10 (6-12 hours post-transfection) and found 20–200 molecules at the filopodial tip38. These levels are consistent with the range that we measure here (Fig. 3A). The shorter time post-transfection clearly affects Myo10 expression levels, which is evident in our data. Fixed cell experiments had higher Myo10 expression overall (48 hours post-transfection) compared to live cell experiments (24 hours post-transfection) (Fig. 1D, E).

Revised version: Prior information on the absolute number of Myo10 molecules in filopodia is sparse. Kerber et al. transfected HeLa cells with low levels of GFP-Myo10 (6-12 hours post-transfection) and found 20–200 molecules at the filopodial tip37. These levels are consistent with the range that we measure here (Figure 3A). The shorter time post-transfection clearly affects Myo10 expression levels, which is evident in our data. Fixed cell experiments had higher Myo10 expression overall (48 hours post-transfection) compared to live cell experiments (24 hours post-transfection) (Figure 1D, E).

Myo10 accretion into filopodial puncta

Earlier version: The distribution of Myo10 among filopodial puncta is log-normal rather than normal (Fig. 3A, B, Supplementary Fig. 4). This stands in contrast to more consistent, normally-distributed quantities of other actin-binding proteins. In fission yeast, proteins in the spindle pole body (e.g., Sad1p), at the cytokinetic contractile ring (e.g., Myo2p, Rho GEF Rgf1p), and within actin patches (e.g., ARPC1, capping protein Acp2p) did not display the wide concentration ranges that we saw for Myo10 in filopodia25. We propose several reasons for the non-Gaussian distribution of Myo10 molecules. Log-normal distributions are commonly associated with growth processes, and the accumulation of Myo10 at filopodial tips fits this description. Filopodia have a large capacity for additional Myo10, and an ample reserve of unactivated Myo10 remains in the cytosol. Slow activation of Myo10 would then allow its travel down filopodia and accretion into puncta. Filopodia operate far from saturation, enabling relatively unrestricted Myo10 accumulation. In contrast, we suspect that the spindle pole body, cytokinetic ring, and actin patches represent cytoskeletal systems constructed from one or more limiting reagents. Limiting reagents put an upper bound on the growth of the system and lead to normal distributions at saturation.

Revised version: The distribution of Myo10 among filopodial puncta is log-normal rather than normal (Figure 3A, B, Figure 3—figure supplement 1). This stands in contrast to more consistent, normally-distributed quantities of other actin-binding proteins. In fission yeast, proteins in the spindle pole body (e.g., Sad1p), at the cytokinetic contractile ring (e.g., Myo2p, Rho GEF Rgf1p), and within actin patches (e.g., ARPC1, capping protein Acp2p) did not display the wide concentration ranges that we saw for Myo10 in filopodia24. We propose several reasons for the non-Gaussian distribution of Myo10 molecules. Log-normal distributions are commonly associated with growth processes, and the accumulation of Myo10 at filopodial tips fits this description. Filopodia have a large capacity for additional Myo10, and an ample reserve of unactivated Myo10 remains in the cytosol. Slow activation of Myo10 would then allow its travel down filopodia and accretion into puncta. Filopodia operate far from saturation, enabling relatively unrestricted Myo10 accumulation. In contrast, we suspect that the spindle pole body, cytokinetic ring, and actin patches represent cytoskeletal systems constructed from one or more limiting reagents. Limiting reagents put an upper bound on the growth of the system and lead to normal distributions at saturation.

Earlier version: Furthermore, the end requirements of cytokinesis and filopodial formation differ. Cytokinesis is a tightly regulated process involving a balance of forces, and the precise timing of each stage has been described in fission yeast47. Disruptions to components of the contractile ring, such as myosin inhibition, affect rates of actin filament disassembly and ring constriction26. In contrast, there is no specific number of filopodia that cells aim to create nor is there an optimal number of Myo10 per filopodium. Therefore, filopodia formation is a much more permissive process than cytokinesis, which could also explain the variation in values we observe of Myo10 compared to contractile ring-associated proteins.

Revised version: Furthermore, the end requirements of cytokinesis and filopodial formation differ. Cytokinesis is a tightly regulated process involving a balance of forces, and the precise timing of each stage has been described in fission yeast46. Disruptions to components of the contractile ring, such as myosin inhibition, affect rates of actin filament disassembly and ring constriction25. In contrast, there is no specific number of filopodia that cells aim to create nor is there an optimal number of Myo10 per filopodium. Therefore, filopodia formation is a much more permissive process than cytokinesis, which could also explain the variation in values we observe of Myo10 compared to contractile ring-associated proteins.

Non-uniform Myo10 filopodial distribution

Earlier version: Myo10 tends to accumulate in filopodia that are at opposite ends of the cell (Fig. 2E, F). Thus, although Myo10 slowly accretes in individual filopodia, the accretion rate also varies spatially. We hypothesize a few different explanations for why Myo10 is not equally distributed around the cell. Potentially, Myo10 is funneled into filopodia that are moving towards specific environmental stimuli that simultaneously affect Myo1048,49. For example, local generation of PIP3 might serve to dock and activate Myo10 in focused regions of the plasma membrane, such as the basolateral surfaces of polarized MDCK cells50,23. The docked Myo10 can then undergo 2D diffusion until it encounters the base of a filopodium. At the base, Myo10 switches to directed movement where it walks processively along the actin bundle to the tip. This overall diffusion and capture process for Myo10 has been observed at the single molecule level in cells, where it has been described as a 3D to 2D to 1D reduction in dimensionality51.

Revised version: Myo10 tends to accumulate in filopodia that are at opposite sides of the cell. These high Myo10 sectors have a higher density of filopodia but a similar quantity of Myo10 per filopodium (Figure 2E, F). We hypothesize a few different explanations for why Myo10 and filopodia not equally distributed around the cell. Potentially, Myo10 is funneled into filopodia that are moving towards specific environmental stimuli that simultaneously affect Myo1047,48. For example, local generation of PIP3 might serve to dock and activate Myo10 in focused regions of the plasma membrane, such as the basolateral surfaces of polarized MDCK cells49,22. The docked but diffusing Myo10 is then more likely to initiate and elongate new filopodium. This overall diffusion and capture process for Myo10 has been observed at the single molecule level in cells, where it has been described as a 3D to 2D to 1D reduction in dimensionality50.

Earlier version: Alternatively, Myo10 could be responding to information encoded in the actin network itself (e.g., nearby stress fibers) or other signaling molecules present in certain filopodia. Myosins can differentiate the structural and chemical properties of actin filaments, including age, tension, curvature, post-translational modifications, and multifilament architecture52. Thus, actin filament networks could be partly responsible for choreographing Myo10 trajectories52. Future studies could identify potential guidance cues that Myo10-positive filopodia are responding to and identify other signaling proteins enriched in filopodia with high Myo10 signal.

Revised version: Alternatively, Myo10 could be responding to information encoded in the actin network itself (e.g., nearby stress fibers) or other local signaling molecules. Myosins can differentiate the structural and chemical properties of actin filaments, including age, tension, curvature, post-translational modifications, and multifilament architecture51. Thus, actin filament networks could be partly responsible for choreographing Myo10 trajectories51. Future studies could identify potential guidance cues that trigger the initiation of Myo10-positive filopodia.

An excess of Myo10 at the filopodial tip

Earlier version: A major finding of this work is that there is often an excess of Myo10 at filopodial tips over the available actin binding-sites for myosin (Fig. 3E-H). Interestingly, there is a mechanism that continues to push Myo10 into the filopodia, even when there may not be available actin to bind. This observation points to two notable features of motility in filopodial systems. The first is that there is no mechanism for a filopodium to shut off entry at the base when the tip is already overfilled with Myo10. The two sites are functionally independent of each other, although retracting filopodia differ in this regard (see below). The second is that Myo10 does not pack in an orderly fashion at the filopodial tip. Although our observations are in the context of an overexpression system, we note that immunofluorescence staining of endogenous Myo10 also shows occasional high-intensity, bulbous filopodial tips20. Our concentration estimates at the filopodial tip suggest that there is sufficient plasma membrane area for the Myo10. Therefore, even though Myo10 disengages from actin, it may remain engaged to the plasma membrane through its PH domains (Fig. 3F, Supplemental Fig. 6). As Myo10 molecules diffuse out of the concentrated region at the tip, they can encounter open sites on the actin filament bundle and reengage, walking back to the tip.

Revised version: A major finding of this work is that there is often an excess of Myo10 at filopodial tips over the available actin binding-sites for myosin (Figure 3E-H). Interestingly, there is a mechanism that continues to push Myo10 into the filopodia, even when there may not be available actin to bind. This observation points to two notable features of motility in filopodial systems. The first is that there is no mechanism for a filopodium to shut off entry at the base when the tip is already overfilled with Myo10. The two sites are functionally independent of each other, although retracting filopodia differ in this regard (see below). The second is that Myo10 does not pack in an orderly fashion at the filopodial tip. Although our observations are in the context of an overexpression system, we note that immunofluorescence staining of endogenous Myo10 also shows occasional high-intensity, bulbous filopodial tips19. Our concentration estimates at the filopodial tip suggest that there is sufficient plasma membrane area for the Myo10. Therefore, even though Myo10 disengages from actin, it may remain engaged to the plasma membrane through its PH domains (Figure 3F, Figure 3—figure supplement 3). As Myo10 molecules diffuse out of the concentrated region at the tip, they can encounter open sites on the actin filament bundle and reengage, walking back to the tip.

Earlier version: The structural model of the filopodium discussed above assumes an orderly, hexagonally close-packed bundle structure all the way to the filopodial tip. Consistent with this picture, EM images of mouse melanoma B16F1 cells indicate bundled parallel actin filaments at the filopodial tip53. However, this is not always the case. In Dictyostelium discoideum cells, actin filaments in the filopodia tip can terminate in free ends and/or appear fragmented, by which short actin filaments are arranged into a non-parallel array54. Bent filopodial tips in mammalian cells overexpressing GFP-Myo10 can display bulbous membrane extensions containing splayed actin filaments or actin filaments arranged in loops55. These fragmented, frayed, or otherwise disorganized ends of the filopodial actin bundle might support more actin-bound Myo10 (Fig. 3H). Indeed, an excess of Myo10 at the tip would compete with actin crosslinkers and help to support fraying. Further investigation is required to understand how Myo10 traffic jams and actin fraying affect the dynamics of actin itself and the activity of its associated proteins.

Revised version: The structural model of the filopodium discussed above assumes an orderly, hexagonally close-packed bundle structure all the way to the filopodial tip. Consistent with this picture, EM images of mouse melanoma B16F1 cells indicate bundled parallel actin filaments at the filopodial tip52. However, this is not always the case. In Dictyostelium discoideum cells, actin filaments in the filopodia tip can terminate in free ends and/or appear fragmented, by which short actin filaments are arranged into a non-parallel array53. Bent filopodial tips in mammalian cells overexpressing GFP-Myo10 can display bulbous membrane extensions containing splayed actin filaments or actin filaments arranged in loops54. These fragmented, frayed, or otherwise disorganized ends of the filopodial actin bundle might support more actin-bound Myo10 (Figure 3H). Indeed, an excess of Myo10 at the tip would compete with actin crosslinkers and help to support fraying. Further investigation is required to understand how Myo10 traffic jams and actin fraying affect the dynamics of actin itself and the activity of its associated proteins.

Number of Myo10 molecules needed for filopodial initiation and second-phase elongation

Earlier version: We observe a median of 160 Myo10 molecules and a minimum of 52 molecules at filopodial initiation (Fig. 4A). These quantities correspond to approximately 5 and 2 molecules per actin filament at a nascent filopodium, respectively. Watanabe et al. also saw a rapid (2-20 s) accumulation of Myo10 recruitment at filopodial initiation19. Myo10 participates in bridging actin filaments in emergent fascin-actin bundles, the actin structures that it prefers3,37,56. Such bridging could help to organize or orient the bundles to allow protrusion, much like the development of the lambda-precursors observed by Svitkina et al53.

Revised version: We observe a median of 160 Myo10 molecules and a minimum of 52 molecules at filopodial initiation (Figure 4A). These quantities correspond to approximately 5 and 2 molecules per actin filament at a nascent filopodium, respectively. Watanabe et al. also saw a rapid (2-20 s) accumulation of Myo10 recruitment at filopodial initiation18. Myo10 participates in bridging actin filaments in emergent fascin-actin bundles, the actin structures that it prefers3,36,55. Such bridging could help to organize or orient the bundles to allow protrusion, much like the development of the lambda-precursors observed by Svitkina et al52.

Earlier version: Second-phase filopodial elongation involves more molecules of Myo10, a median of 290 (Fig. 4B). This increase in number likely reflects the larger available pool of active Myo10 present at an established filopodial tip. Because only 160 molecules are needed for filopodial initiation, we believe that Myo10 quantities do not limit second-phase extension if initiation and second-phase elongation have similar mechanical requirements. Focal adhesion proteins, such as vinculin and integrin-β1, have been previously identified as essential components during the second-phase of filopodia extension15. Other actin-associated proteins (e.g., Arp2/3 and vinculin) are also present at this stage15. The duration required for focal adhesion proteins to translocate into a filopodium, coupled with actin crosslinking and bundling in the new filopodium, results in an increased accumulation of Myo10 at the local level. This accumulation is a consequence of Myo10 awaiting interaction with other protein partners at the filopodium tip. We speculate that the entire process takes longer than at the cell body plasma membrane, where the concentration of available focal adhesion proteins is higher.

Revised version: Second-phase filopodial elongation involves more molecules of Myo10, a median of 290 (Figure 4B). This increase in number likely reflects the larger available pool of active Myo10 present at an established filopodial tip. Because only 160 molecules are needed for filopodial initiation, we believe that Myo10 quantities do not limit second-phase extension if initiation and second-phase elongation have similar mechanical requirements. Focal adhesion proteins, such as vinculin and integrin-β1, have been previously identified as essential components during the second-phase of filopodia extension14. Other actin-associated proteins (e.g., Arp2/3 and vinculin) are also present at this stage14. The duration required for focal adhesion proteins to translocate into a filopodium, coupled with actin crosslinking and bundling in the new filopodium, results in an increased accumulation of Myo10 at the local level. This accumulation is a consequence of Myo10 awaiting interaction with other protein partners at the filopodium tip. We speculate that the entire process takes longer than at the cell body plasma membrane, where the concentration of available focal adhesion proteins is higher.

Myo10 entry into filopodia switches off upon retraction

Earlier version: Myo10 accumulates within puncta after filopodial initiation and second-phase elongation (Fig. 4D, E), adding several hundred molecules on the minutes timescale. However, in retracting filopodia, accretion of additional Myo10 ceases (Fig. 4F). These retracting filopodia frequently have immediate neighbors that are not retracting. Therefore, the mechanism that switches off Myo10 traffic must be highly focused to single filopodia. Moreover, this switching mechanism must operate at the base of the filopodium where the Myo10 traffic would otherwise enter. Two mechanisms seem plausible given these considerations. Either the retracting filopodium generates a Myo10-inhibitory signal that diffuses out and rapidly dissipates, or the actin architecture of the filopodium is altered and blocks Myo10 motility. These findings offer a quantitation of the dynamic behavior of Myo10 during retrograde flow, but further research is necessary to uncover the mechanisms governing switching of filopodial movement.

Revised version: Myo10 accumulates within puncta after filopodial initiation and second-phase elongation (Figure 4D, E), adding several hundred molecules on the minutes timescale. However, in retracting filopodia, accretion of additional Myo10 ceases (Figure 4F). These retracting filopodia frequently have immediate neighbors that are not retracting. Therefore, the mechanism that switches off Myo10 traffic must be highly focused to single filopodia. Moreover, this switching mechanism must operate at the base of the filopodium where the Myo10 traffic would otherwise enter. Two mechanisms seem plausible given these considerations. Either the retracting filopodium generates a Myo10-inhibitory signal that diffuses out and rapidly dissipates, or the actin architecture of the filopodium is altered and blocks Myo10 motility. These findings offer a quantitation of the dynamic behavior of Myo10 during retrograde flow, but further research is necessary to uncover the mechanisms governing switching of filopodial movement.

Multiple motors interfere to slow dynamic movements at the filopodial tip

Earlier version: The presence of increasing quantities of Myo10 slow filopodial movements in all three phases of the filopodial lifecycle (Fig. 4G-I). In reconstituted motility systems with multiple molecules, mechanically coupled motors typically move more slowly than uncoupled (or single) motors. This slowing is attributed to strain-sensitive motor coupling and motor dissociation limiting the collective velocity to some degree57,58. We propose that similar drag mechanisms operate here to slow the dynamics of filopodia.

Revised version: The presence of increasing quantities of Myo10 slow filopodial movements in all three phases of the filopodial lifecycle (Figure 4G-I). In reconstituted motility systems with multiple molecules, mechanically coupled motors typically move more slowly than uncoupled (or single) motors. This slowing is attributed to strain-sensitive motor coupling and motor dissociation limiting the collective velocity to some degree56,57. We propose that similar drag mechanisms operate here to slow the dynamics of filopodia.

Earlier version: Recent studies address the influence of filopodial-tip directed myosin motors on filopodial elongation. In work by Cirilo et al., wildtype and fast-mutant Myo3 motors also localize to filopodial tips, and the filopodial elongation rate correlates with the Myo3 speed21. Myo3 is a slower motor than Myo10 (70 nm/s vs. 300-600 nm/s), and the Myo3-decorated filopodia move more slowly than we typically observe here21. These observations with slower myosins are consistent with our drag-based interference of filopodial dynamics: a coupled collection of slower motors would produce more drag than a collection of faster motors, all else being equal.

Revised version: Recent studies address the influence of filopodial-tip directed myosin motors on filopodial elongation. In work by Cirilo et al., wildtype and fast-mutant Myo3 motors also localize to filopodial tips, and the filopodial elongation rate correlates with the Myo3 speed20. Myo3 is a slower motor than Myo10 (70 nm/s vs. 300-600 nm/s), and the Myo3-decorated filopodia move more slowly than we typically observe here20. These observations with slower myosins are consistent with our drag-based interference of filopodial dynamics: a coupled collection of slower motors would produce more drag than a collection of faster motors, all else being equal.

Limitations

Earlier version: Our experiments were conducted with the intention of minimizing biases and errors, but our results do come with caveats. Our results are within the context of an overexpression system, and other systems expressing endogenous Myo10 may express at a lower overall level than here. However, cells expressing endogenous Myo10 also show prominent dense tip-localization by immunofluorescence20, so the Myo10 organization we see here does exist physiologically. Our goal in this study was to understand Myo10’s role in filopodial initiation and dynamics in a defined system where Myo10 is the limiting factor. U2OS cells express undetectable levels of endogenous Myo10 (Supplementary Fig. 1E), so we are indeed observing the influence of Myo10 in a situation where Myo10 is the limiting factor for filopodia formation.

Revised version: Our experiments were conducted with the intention of minimizing biases and errors, but our results do come with caveats. Our results are within the context of an overexpression system, and other systems expressing endogenous Myo10 may express at a lower overall level than here. However, cells expressing endogenous Myo10 also show prominent dense tip-localization by immunofluorescence19, so the Myo10 organization we see here does exist physiologically. Our goal in this study was to understand Myo10’s role in filopodial initiation and dynamics in a defined system where Myo10 is the limiting factor. U2OS cells express undetectable levels of endogenous Myo10 (Figure 1—figure supplement 1E), so we are indeed observing the influence of Myo10 in a situation where Myo10 is the limiting factor for filopodia formation.

Regarding aspects of microscopy and automated image processing, very dim Myo10 spots could have escaped detection. During fixed and live imaging, we are unable to detect single molecules of Myo10 because we do not have the sensitivity of TIRF. Furthermore, our estimates of molecules are predicated on the calibration curve of the HaloTag standard protein on the SDS-PAGE gels, which is likely the highest source of error on our molecule counts. Despite these concerns, our values still provide a sense of the magnitude of Myo10 molecules within cellular structures.

Ideas & Speculation

Earlier version: One of our key findings is the excess of Myo10 compared to available actin at filopodial tips. We propose that there is a sizable population of Myo10 docked on the plasma membrane but not to actin (Figure. 4G). Alternatively, actin filaments that fray from the fascin-actin bundle yield more binding sites for Myo10 (Fig. 4H). Myo10 moves processively along actin bundles, even disorganized ones formed by crowding reagents37. Any single overloaded filopodium may use either or both mechanisms.

Revised version: One of our key findings is the excess of Myo10 compared to available actin at filopodial tips. We propose that there is a sizable population of Myo10 docked on the plasma membrane but not to actin (Figure. 4G). Alternatively, actin filaments that fray from the fascin-actin bundle yield more binding sites for Myo10 (Figure 4H). Myo10 moves processively along actin bundles, even disorganized ones formed by crowding reagents36. Any single overloaded filopodium may use either or both mechanisms.

We argue that there is an interplay between incoming Myo10 that release upon encountering a traffic jam of other Myo10 molecules, and Myo10 molecules that are part of the traffic jam that are “pushed off” (i.e., naturally dissociate and are replaced) by new arrivals. Myo10 cargoes may need to handle either the actin-attached or actin-detached situation, which has implications for force-sensing cargoes such as integrins.

Earlier version: At the filopodial tip, we expect that newly formed actin will be rapidly occupied by Myo10 due to the high quantity on standby. However, fascin must eventually bind to facilitate the bundling of new actin. While Myo10 cannot bind to the interior of a fascin-actin bundle due to insufficient space, the converse is also true: fascin cannot bundle a fully Myo10-bound actin filament until the Myo10 dissociates. Once there is an opening, fascin can outcompete Myo10 due to the high avidity interaction between fascin and actin59. Therefore, the dynamic association/dissociation of Myo10 at the filopodial tip may be essential for the continued elongation of the filopodium.

Revised version: At the filopodial tip, we expect that newly formed actin will be rapidly occupied by Myo10 due to the high quantity on standby. However, fascin must eventually bind to facilitate the bundling of new actin polymerized at the filopodial tip. While Myo10 cannot bind to the interior of a fascin-actin bundle due to insufficient space, the converse is also true: fascin cannot crosslink at a Myo10-bound actin filament site until the Myo10 dissociates. However, this Myo10 dissociation occurs frequently as part of the stepping process. Once there is an opening, fascin can outcompete Myo10 due to the high avidity interaction between fascin and actin.

Added in the revised version: Fascin binding is essentially irreversible, as we saw no disassembly of a suspended fascin-actin bundle even after 1 h of continuous buffer wash58. Therefore, the dynamic association and dissociation of Myo10 at the filopodial tip with subsequent, irreversible fascin binding may be essential for the continued elongation of the filopodium.

Conclusion

Knowing the number of Myo10 molecules in a filopodium provides insight into molecular packing geometries in confined intracellular spaces. This study showcases the high density of molecules that can be packed into tight cellular compartments, captures the number of Myo10 molecules needed for filopodial initiation, and finds that Myo10 entry into filopodia switches off upon filopodial retraction. Our results contribute to understandings of molecular stoichiometries in filopodia and Myo10 abundance in these structures. Our protocol also provides a framework for future studies examining the factors that tune Myo10 density and distribution.

Methods

Plasmids and constructs

Earlier version: Full-length human Myo10 sequence was constructed in a pTT5 vector backbone plasmid by Gibson assembly60. The construct includes an N-terminal Halotag (Promega; GenBank: JF920304.1), the full-length human Myo10 sequence (nucleotide sequence from GenBank: BC137168.1), and C-terminal Flag tag (GATTATAAAGATGATGATGATAAA). A complete DNA sequence of the insert is included in supplemental information.

Revised version: Full-length human Myo10 sequence was constructed in a pTT5 vector backbone plasmid by Gibson assembly59. The construct includes an N-terminal Halotag (Promega; GenBank: JF920304.1), the full-length human Myo10 sequence (nucleotide sequence from GenBank: BC137168.1), and C-terminal Flag tag (GATTATAAAGATGATGATGATAAA). A complete DNA sequence of the insert is included in supplemental information.

Cell culture

Earlier version: U2OS cells were passaged every 2 days and used under passage number 10 after thawing. Cells were grown in Gibco 1x DMEM (ThermoFisher, 11995073) supplemented with 10% fetal bovine serum at 37 °C and 5% CO2.

Revised version: U2OS cells (ATCC HTB-96, tested negative for mycoplasma by PCR and DAPI staining) were passaged every 2 days and used under passage number 10 after thawing. Cells were grown in Gibco 1x DMEM (ThermoFisher 11995073) supplemented with 10% fetal bovine serum at 37 °C and 5% CO2.

Transient transfection and coverslip seeding for fixed cell imaging

Earlier version: Cells were transiently transfected with 1 µg of the HaloTag-Myo10-FlagTag plasmid and 0.2 µg of a calmodulin plasmid using FuGENE HD Transfection reagent. Forty-eight hours after transfection, cells were detached using Accutase and seeded onto Ibidi 8-well chamber slides coated with 20 µg/ml laminin for imaging and onto 6 well dishes for SDS-PAGE analysis. After 3-4 hours, cells were collected for SDS-PAGE analysis at the same time as cells were fixed for microscopy. Three bioreplicates were conducted.

Revised version: Cells were transiently transfected with 1 µg of the HaloTag-Myo10-FlagTag plasmid and 0.2 µg of a calmodulin plasmid using FuGENE HD Transfection reagent (Promega E2311). Forty-eight hours after transfection, cells were detached using Accutase (Corning MT25058CI) and seeded onto 8-well chamber slides (Ibidi 80807) coated with 20 µg/ml laminin (Sigma-Aldrich CC095-M) for imaging and onto 6 well dishes for SDS-PAGE analysis. After 3-4 hours, cells were collected for SDS-PAGE analysis at the same time as cells were fixed for microscopy. Three bioreplicates were conducted.

Transient transfection and coverslip seeding for live cell imaging

Earlier version: Cells were transiently transfected with 1 µg of the HaloTag-Myo10-FlagTag plasmid and 0.2 µg of a calmodulin plasmid using Lipofectamine 2000. Twenty-four hours after transfection, cells were seeded onto #1.5 coverglass bottom 35 mm petri dishes (Cellvis or MatTek) coated with 20 µg/ml laminin for imaging and onto 6 well dishes for SDS-PAGE analysis. After 3-4 hours, cells were collected for SDS-PAGE analysis right before cells were live-imaged. Three bioreplicates were conducted.

Revised version: Cells were transiently transfected with 1 µg of the HaloTag-Myo10-FlagTag plasmid and 0.2 µg of a calmodulin plasmid using Lipofectamine 2000 (Invitrogen 11668-027). Twenty-four hours after transfection, cells were seeded onto #1.5 coverglass bottom 35 mm petri dishes (Cellvis D35-10-1.5-N or MatTek P35G-1.5-14-C) coated with 20 µg/ml laminin for imaging and onto 6 well dishes for SDS-PAGE analysis. After 3-4 hours, cells were collected for SDS-PAGE analysis right before cells were live-imaged. Three bioreplicates were conducted.

Preparing cell lysates for SDS-PAGE

Earlier version: Cells growing in the 6 well plate were first incubated with Accutase for 15 min at room temperature. For live cell experiments, the Accutase additionally contained 0.75 µM TMR-HaloLigand. Accutase was neutralized with DMEM +10 % FBS and the cells were collected with a 3:30 min spin at 400 rcf at room temperature. The pellet was resuspended in 100 µl cell media, and 10 µl of the cells were removed and mixed with 1x Trypan blue for cell counting. The remaining cells were combined with 400 µl cell media and subjected to a 5 min, 400 rcf spin at room temperature. Cells were moved to ice post-spin and lysed in RIPA buffer containing 1 mM PMSF. Lysate samples were stored at -80 °C.

Revised version: Cells growing in the 6 well plate were first incubated with Accutase for 15 min at room temperature. For live cell experiments, the Accutase additionally contained 0.75 µM TMR-HaloLigand (Promega G8251). Accutase was neutralized with DMEM +10 % FBS and the cells were collected with a 3:30 min spin at 400 rcf at room temperature. The pellet was resuspended in 100 µl cell media, and 10 µl of the cells were removed and mixed with 1x Trypan blue for cell counting. The remaining cells were combined with 400 µl cell media and subjected to a 5 min, 400 rcf spin at room temperature. Cells were moved to ice post-spin and lysed in RIPA buffer containing 1 mM PMSF. Lysate samples were stored at −80 °C.

SDS gel analysis for Myo10 quantitation

HaloTag standard protein (Promega G4491) samples, with 0.5 µM TMR-HaloLigand, were prepared for final masses of 1.25 ng, 2.5 ng, 5 ng, 10 ng and 15 ng and cell lysate samples (from live and fixed cell experiments) containing 50,000 cells on the gel. Excess TMR-HaloLigand was incubated for 10 min on ice with only lysates from the fixed cell experiments. All samples were supplemented with an SDS loading buffer (0.1 M DTT, 2% SDS, 0.05% bromophenol blue, 0.05M Tris-HCl, 10% glycerol, pH 6.8) and heated at 70°C for 10 min before loading onto a 4-20% Mini-PROTEAN TGX Stain-Free protein gel (BioRad 4568095). The gel was run at 180 V for 45 min and imaged on a ChemiDoc. Images of the stain-free, AF647, and rhodamine channels were taken. We used the Gel Analysis plugin in ImageJ to compare signal intensity of the gel bands. A standard curve was generated from the HaloTag standard protein TMR signal to estimate the number of Myo10 molecules per total transfected cells loaded (total cells loaded times the fraction transfected). Microscopy was used to count the percentage of transfected cells from ∼105-190 randomly surveyed cells per bioreplicate.

Western blotting for Myo10 in U2OS cells

Earlier version: We loaded 50,000 cells of wild-type and Myo10-transfected U2OS cells for SDS gel analysis as described earlier. The gel was transferred to PVDF membrane using a Pierce Power Blotter semi-dry electrophoretic transfer cell (1.3 A constant for 12 min). The blot was blocked with 5% milk in 1x TBST, 1 hour shaking at room temperature. For 1 hour at room temperature, the blot was incubated with anti-Myosin10 (Novus NBP1-87748) and anti-{3-tubulin (Invitrogen 22833) antibodies in 5% milk in 1x TBST. The blot was rinsed with 5% milk in 1x TBST for 10 min, three total washes. Then, the blot was incubated with anti-Rabbit-HRP (Cell Signaling 7074) and anti-Mouse-HRP (Cell Signaling 7076) antibodies in 1x TBST, 1 hour shaking at room temperature. After three 10 min washes with 1x TBST, chemiluminescent substrate (Thermo Scientific, SuperSignal West Femto, 34094) was applied to the blot for 2 min and subsequently imaged on a Chemidoc.

Revised version: We loaded 50,000 cells of wild-type and Myo10-transfected U2OS cells for SDS gel analysis as described earlier. The gel was transferred to PVDF membrane using a Pierce Power Blotter semi-dry electrophoretic transfer cell (1.3 A constant for 12 min). The blot was blocked with 5% milk in 1x TBST, 1 hour shaking at room temperature. For 1 hour at room temperature, the blot was incubated with anti-Myosin10 (Novus NBP1-87748) and anti-,B-tubulin (Invitrogen 22833) antibodies in 5% milk in 1x TBST. The blot was rinsed with 5% milk in 1x TBST for 10 min, three total washes. Then, the blot was incubated with anti-Rabbit-HRP (Cell Signaling 7074) and anti-Mouse-HRP (Cell Signaling 7076) antibodies in 1x TBST, 1 hour shaking at room temperature. After three 10 min washes with 1x TBST, chemiluminescent substrate (Thermo Scientific, SuperSignal West Femto, 34094) was applied to the blot for 2 min and subsequently imaged on a Chemidoc.

TMR-HaloLigand labeling efficiency

To test TMR-HaloLigand labeling efficiency, 6.8 µM of HaloTag standard protein was labeled with TMR-HaloLigand in 5.8x molar excess in a 20 µl reaction volume, on ice for 10 min. Bio-Beads SM-2 (BioRad 152-8920) were washed with methanol 3x, distilled water 3x, and PBS 3x before they were added to ¼ volume of the reaction. After 1 hour and 40 min at RT in the dark, the reaction tube was spun for 3 min, 600 rcf spin. The supernatant was recovered and the absorbance spectrum was measured on Nanodrop. Absorbances at 280 nm and 553 nm were used for protein and dye concentrations, respectively. To account for protein that may have nonspecifically bound to the beads, pre-bead TMR-HaloLigand-labeled HaloTag standard protein was loaded on an SDS-PAGE gel alongside the post-bead supernatant. Comparing the signal intensities of pre- and post-bead gel bands allowed for a corrected post-beads protein concentration value.

Ligand labeling saturation for microscopy

Earlier version: To test ligand labeling saturation for microscopy, cells were fixed for 20 min in a solution comprising: 4% PFA, 0.08% Triton, 1:1000 DAPI in PEM buffer and a series of TMR-HaloLigand concentrations (in µM: 0.05, 0.1, 0.5, 1, 2.5, 5). Live cells were incubated with DMEM +10 % FBS containing TMR-HaloLigand (in µM: 0.05, 0.1, 0.5, 0.75, 1, 2.5 µM) at 37°C, 5% CO2 in the dark. After 10 min, cells were washed and incubated with GOC during live imaging as described earlier.

Revised version: To test ligand labeling saturation for microscopy, cells were fixed for 20 min in a solution comprising: 4% PFA, 0.08% Triton, 1:1000 DAPI in PEM buffer and a series of TMR-HaloLigand concentrations (in µM: 0.05, 0.1, 0.5, 1, 2.5, 5). Live cells were incubated with DMEM +10 % FBS containing TMR-HaloLigand (in µM: 0.05, 0.1, 0.5, 0.75, 1, 2.5 µM) at 37°C, 5% CO2 in the dark. After 10 min and immediately prior to imaging, cells were washed and incubated with DMEM + 10 % FBS + GOC (an enzymatic oxygen-scavenger system: 4.5 mg/ml glucose, 0.5% β-mercaptoethanol, 4.3 mg/ml glucose oxidase, and 0.7 mg/ml catalase).

Image analysis was done using in-house Python scripts. Images were background subtracted using a rolling ball radius of 50 pixels in ImageJ. To obtain total intracellular Myo10 signal in fixed cells, the phalloidin stain was used as a guide to manually draw bound cells. Filopodial puncta, segmented by watershed segmentation, were used to generate a convex hull mask. Signal inside the mask was summed.

To obtain total intracellular Myo10 signal in live cells, filopodial puncta, segmented by watershed segmentation, were used to generate a convex hull mask. Signal inside the convex hull mask was summed. Because there was higher, uneven background that persisted after the rolling ball subtraction in higher TMR-HaloLigand live samples, segmented Myo10 puncta were summed and plotted as well. Plotting only segmented puncta appeared to reduce the contribution of background pixels inside the convex hull mask that could falsely inflate final Myo10 intracellular signal.

Fixed-cell fluorescence microscopy

Earlier version: For Myo10 fixed cell imaging, cells were fixed for 20 min in a solution comprising: 4% PFA, 0.08% Triton, 2.5 µM TMR-HaloLigand, and 1:1000 DAPI in PEM buffer. After 3 PBS washes (4 min each), cells were stained with 1 mM Phalloidin-AF633 in 1% BSA for 20 min. Cells were washed 3x with PBS (4 min each) before immediate imaging. Fluorescence images were taken on an Axiovert using a 60x water objective and a gain of 125 and 175 ms exposure (below detector saturation). To prevent photobleaching, cells were kept in the dark and imaged immediately upon illumination of the selected field of view. Samples were imaged within a day of preparation. TMR was visualized using green light and AF633 was visualized with red light.

Revised version: For Myo10 fixed cell imaging, cells were fixed for 20 min in a solution comprising: 4% PFA, 0.08% Triton, 2.5 µM TMR-HaloLigand, and 1:1000 DAPI in PEM buffer. After 3 PBS washes (4 min each), cells were stained with 1 mM Phalloidin-AF647 (Invitrogen A22287) in 1% BSA for 20 min. Cells were washed 3x with PBS (4 min each) before immediate imaging. Fluorescence images were taken on an Axiovert using a 60x water objective and a gain of 125 and 175 ms exposure (below detector saturation). To prevent photobleaching, cells were kept in the dark and imaged immediately upon illumination of the selected field of view. Samples were imaged within a day of preparation. TMR was visualized using green light and AF633 was visualized with red light.

Analyzing fixed cell images

Earlier version: Image analysis was done using in-house Python scripts. To calculate the number of Myo10 molecules in the fluorescence images, the Myo10-stained images were first background subtracted: the average signal of a 56 x 56 pixel square near the cell body was calculated and subtracted from each pixel of the TIFF image. Next, the phalloidin-stained cell image and Myo10-stained image were filtered to remove non-cell objects and subjected to watershed segmentation. To generate a ‘cell body mask’, an erosion function followed by an opening function were applied to the phalloidin-stained cell image. The Myo10 image mask was defined as the ‘total cell mask.’ The ‘cell body mask’ was subtracted from the ‘total cell mask’ for a ‘filopodia mask.’ Connected component analysis (CCA) was performed on the ‘filopodia mask’ to obtain masks for all filopodia-localized Myo10 puncta. For the CCA, pixels were considered neighbors if they were connected through a maximum of 2 orthogonal hops. Segmented puncta (i.e. the connected regions) were then inspected in Napari61. Puncta belonging to the same filopodium were manually combined into the same filopodium mask. Signal in each filopodium mask and within the ‘cell body mask’ were summed. For each bioreplicate, a ‘total pool’ of Myo10 molecules was calculated by multiplying the average Myo10 molecules per transfected cell (as determined from the SDS-PAGE gel) by the number of cell images analyzed. Within each bioreplicate, the signal intensities were summed for all cells. Therefore, to get the number of molecules within a Myo10 filopodium, the filopodium intensity was divided by the bioreplicate signal intensity sum and multiplied by ‘total pool.’ To calculate percentage of Myo10-positive filopodia, total filopodia per cell were manually counted.

Revised version: Image analysis was done using in-house Python scripts. To calculate the number of Myo10 molecules in the fluorescence images, the Myo10-stained images were first background subtracted: the average signal of a 56 x 56 pixel square near the cell body was calculated and subtracted from each pixel of the TIFF image. Next, the phalloidin-stained cell image and Myo10-stained image were filtered to remove non-cell objects and subjected to watershed segmentation. To generate a ‘cell body mask’, an erosion function followed by an opening function were applied to the phalloidin-stained cell image. The Myo10 image mask was defined as the ‘total cell mask.’ The ‘cell body mask’ was subtracted from the ‘total cell mask’ for a ‘filopodia mask.’ Connected component analysis (CCA) was performed on the ‘filopodia mask’ to obtain masks for all filopodia-localized Myo10 puncta. For the CCA, pixels were considered neighbors if they were connected through a maximum of 2 orthogonal hops. Segmented puncta (i.e. the connected regions) were then inspected in Napari60. Puncta belonging to the same filopodium were manually combined into the same filopodium mask.

Added in the revised version: To convert fluorescence intensities by microscopy into the number of Myo10 molecules, let In represent the sum of the fluorescence signal for n = ∼50 cell images in a bioreplicate (each background corrected, and each integrated over their ‘total cell masks’). This fluorescence signal arises from n m molecules, where ⟨m⟩ denotes the expected number of Myo10 molecules per transfected cell as determined by the SDS-PAGE analysis. Let IROI represent the background-corrected and summed intensity of a single region-of-interest (e.g., a filopodial punctum, or one cell). The number of Myo10 molecules found within that region of interest, r, is then:

Added in the revised version: This procedure was repeated using the ⟨m⟩, n, and In values for each bioreplicate (see Figure 1— figure supplement 3 for a graphical representation of the math). To calculate percentage of Myo10-positive filopodia, total filopodia per cell were manually counted.

Measuring Myo10 orientation distribution in fixed cell filopodia

Earlier version: To generate the rose plots of the Myo10 distribution in filopodia, each cell was divided into 20 radial sections from the cell’s center of mass. The Myo10 filopodial molecules were averaged within each section. The section with the highest molecules was aligned to degree 0 for each cell’s rose plot. Total count of Myo10 puncta analyzed per section was also calculated. Puncta were not manually combined if belonging to the same filopodium. Therefore, puncta that were at the border of two radial sections were counted once for each section. For the rose plots in panel 2D, a randomly selected empty Myo10 section was aligned to degree 0 if more than one section contained no Myo10. The standard error bars represent the standard error of the mean molecules per section after 500 iterations of bootstrapping. In each iteration, 150 cells were randomly selected with replacement.

Revised version: To generate the rose plots of the Myo10 distribution in filopodia, each cell was divided into 20 radial sections from the cell’s center of mass. The Myo10 filopodial molecules were averaged within each section. The section with the highest molecules was aligned to degree 0 for each cell’s rose plot. Total count of Myo10 puncta analyzed per section was also calculated.

Added in the revised version: Puncta were not manually combined if belonging to the same filopodium. Therefore, puncta that were at the border of two radial sections were counted once for each section. For the rose plots in panel 2D, a randomly selected empty Myo10 section was aligned to degree 0 if more than one section contained no Myo10. The standard error bars represent the standard error of the mean molecules per section after 500 iterations of bootstrapping. In each iteration, 150 cells were randomly selected with replacement.

Measuring local Myo10 concentrations at the tips of fixed cell filopodia

Earlier version: To obtain estimates of local Myo10 concentrations in filopodia tips, 10 Myo10 tip-localized puncta were randomly chosen from 3 different cell images of each bioreplicate set. The length of each punctum was measured in ImageJ. The volume occupied by the Myo10 punctum was estimated using the volume of a cylinder: length = height and width = 2*radius, where radius was assumed to be 100 nm (reported average radius) because of the resolution limit. The signal intensity of the Myo10 punctum was converted to molecules as described above. Available actin monomer concentration for Myo10 binding is described in Supplementary Figure 5. Available membrane surface for Myo10 binding is described in Supplementary Figure 6.

Revised version: To obtain estimates of local Myo10 concentrations in filopodia tips, 10 Myo10 tip-localized puncta were randomly chosen from 3 different cell images of each bioreplicate set. The length of each punctum was measured in ImageJ. The volume occupied by the Myo10 punctum was estimated using the volume of a cylinder: length = height and width = 2*radius, where radius was assumed to be 100 nm (reported average radius) because of the resolution limit. The signal intensity of the Myo10 punctum was converted to molecules as described above.

Added in the revised version: Available actin monomer concentration for Myo10 binding is described in Figure 3— figure supplement 2. Available membrane surface for Myo10 binding is described in Figure 3— figure supplement 3.

Staining cells for live cell imaging

For Myo10 live imaging, cells were incubated with DMEM +10 % FBS containing 0.75 µM TMR-HaloLigand at 37 °C in the dark. After 10 min, cells were quickly washed 2x with DMEM +10 % FBS washes, and one 5 min wash. Immediately before imaging, cells were replaced with DMEM +10 % FBS containing 1x GOC. Live cells were imaged on a 37 °C heated objective using a gain of 100 and 300 ms exposure (below detector saturation). Ten random cells per bioreplicate were recorded for <6 minutes (three bioreplicates in total). For each bioreplicate, additional snapshots of ∼50 living cells were taken to sample the intensity distribution for the fluorescence intensity-to-molecule conversions. Fixed cell samples were prepared in parallel to obtain transfection efficiencies of the bioreplicates.

Analyzing live cell movies

Earlier version: Image analysis was partially done using in-house Python scripts. To get an intensity distribution for the fluorescence intensity-to-molecule conversions, the Myo10-stain snapshots were first background subtracted using a rolling ball radius of 50 pixels in ImageJ. To obtain total intracellular Myo10 signal, filopodial puncta were used as bounding edges for a manually drawn mask. Signal inside the drawn mask was summed. For each bioreplicate, a ‘total pool’ of Myo10 molecules was calculated by multiplying the average Myo10 molecules per transfected cell by the number of cell images analyzed. Within each bioreplicate, the signal intensities were summed for all cells. Therefore, to get the number of molecules within a Myo10 punctum, the punctum intensity was divided by the bioreplicate signal intensity sum and multiplied by the ‘total pool.’ Myo10 puncta in the live cell movies were segmented and tracked using the Fiji plugin TrackMate62. Trackmate parameters used: LoG detector with estimated blob diameter = 1 µm, median filter, sub-pixel localization; LAP tracker with frame to frame linking = 0.5 µm, track segment gap closing = 1 µm (max frame gap=2), track segment splitting = 0.5 µm, track segment merging = 1 µm.

Revised version: Image analysis was partially done using in-house Python scripts. To get an intensity distribution for the fluorescence intensity-to-molecule conversions, the Myo10-stain snapshots were first background subtracted using a rolling ball radius of 50 pixels in ImageJ. To obtain total intracellular Myo10 signal, filopodial puncta were used as bounding edges for a manually drawn mask. Signal inside the drawn mask was summed. Conversion of fluorescence signal intensity to number of molecules was performed as described above for fixed cells. Myo10 puncta in the live cell movies were segmented and tracked using the Fiji plugin TrackMate61. Trackmate parameters used: LoG detector with estimated blob diameter = 1 µm, median filter, sub-pixel localization; LAP tracker with frame to frame linking = 0.5 µm, track segment gap closing = 1 µm (max frame gap=2), track segment splitting = 0.5 µm, track segment merging = 1 µm.

After completing TrackMate analysis, all trajectories in the cell movies were manually inspected for occurrence of three types of filopodial events. The start of the three events are defined by when TrackMate first segments the punctum. “Initiation” is when a Myo10 punctum gathers at the plasma membrane and then shoots into a filopodium. “Second-phase elongation” is a new Myo10 punctum that emerges from an existing one. “Retraction” is when a Myo10 punctum, observed in frame 1 of the movie, begins moving back towards the cell body. The track IDs were recorded for each trajectory type and further inspected. A few puncta that TrackMate segmented were manually removed because they were clearly background noise or incorrectly tracked puncta (e.g., a punctum from another filopodium that briefly crosses paths, or filopodial puncta that were detected by TrackMate only well after initiation). Any selected trajectories containing Myo10 puncta merging or splitting events were excluded from velocity analysis.

Earlier version: For Figure 4 parts A-C, Myo10 puncta molecules in the first two frames of each trajectory were averaged. For Figure 4 parts D-F, the starting punctum molecules are subtracted from subsequent frames of each trajectory. Parts D-E follow the trajectory of each punctum that is initially identified for a certain filopodial event, but then those puncta could have engaged in other activities during their lifetime (e.g. retract). GAM fitting was only applied to a subset of the total data to prevent a few long trajectories from dominating the trend. For Figure 4 parts G-H, velocity analysis does not only account for Myo10 punctum that moved away from the starting point of the trajectory. For Figure 4 parts G-I, puncta merging or splitting events were excluded from analysis.

Revised version: For Figure 4 parts A-C, Myo10 puncta molecules in the first two frames of each trajectory were averaged. For Figure 4 parts D-F, the starting punctum molecules are subtracted from subsequent frames of each trajectory. Parts D-E follow the trajectory of each punctum that is initially identified for a certain filopodial event, but then those puncta could have engaged in other activities during their lifetime (e.g., retract). GAM fitting was only applied to a subset of the total data to prevent a few long trajectories from dominating the trend. For Figure 4 parts G-H, note that velocity analysis includes a few Myo10 puncta that switch direction within a single trajectory (e.g., a retracting punctum that then elongates). For Figure 4 parts G-I, puncta merging or splitting events were excluded from analysis.

Added in the revised version: Statistical analysis

Added in the revised version: Statistical tests were performed using R. The statistical tests used for each quantification are stated in the figure legends and within the main text when each figure is initially referenced.

Added in the revised version: Supplementary figures

Added in the revised version: HaloTag ligand-TMR specifically labels HaloTag-Myo10 and reliably reflects the total Myo10 present in U2OS cells.

Added in the revised version: (A) Epifluorescence image of exogenously expressed HaloTag-Myo10 in U2OS cells. Flag-tag was labeled with monoclonal anti-Flag M2 antibody (Sigma, F1804) and goat anti-mouse AF647 (green). HaloTag was labeled with HaloTag ligand-TMR (magenta). Merged image shows high colocalization between labeling of both tags. Scale bar = 10 µm. (B) In fixed cells: distribution of the percent of total filopodia (manually counted) per cell that were Myo10-positive as determined by quantification of 150 cell images. Min = 46.67%, median = 92.35%, max = 100%. (C) Left SDS-PAGE lanes: full-length image of Fig 1B. 50,000 cells from 6 separate U2OS transient transfections. Bioreplicate 5 has 10,000 cells loaded, indicated by a red asterisk (*). Bioreplicate 1, 2 and 3: for live cell analysis. Bioreplicate 4, 5 and 6: for fixed cell analysis. Stain-free shows total protein signal, and TMR illumination shows only HaloTag-Myo10 signal. Right SDS-PAGE lanes: HaloTag standard and U2OS lysates were imaged on the same gel. Signal was integrated for full-length Myo10 (at ∼250 kDa) and Myo10 aggregated in the wells. Integrations of band intensity was done using ImageJ’s Gel Analysis plug-in. (D) SDS-PAGE gel of wildtype U2OS (WT) labeled with HaloTag ligand-TMR vs HaloTag-Myo10 overexpressed in U2OS cells (M10). 50,000 cells of each sample were loaded onto the same gel. Stain-free shows total protein signal, and TMR illumination shows only HaloTag-Myo10 signal. Negative control WT indicates no nonspecific HaloTag ligand-TMR labeling. (E) Left: SDS-PAGE gel of wildtype U2OS (WT) labeled with HaloTag ligand-TMR vs HaloTag-Myo10 overexpressed in U2OS cells (M10). 50,000 cells of each sample loaded. Right: Immunoblot for Myo10 (NBP1-87748) and P-tubulin (Invitrogen 22833). (F) Epifluorescence image of wildtype U2OS cells labeled with HaloTag ligand-TMR. Negative control microscopy sample indicates no nonspecific HaloTag ligand-TMR labeling.

Added in the revised version: HaloTag ligand-TMR robustly labels HaloTag-Myo10.

Added in the revised version: (A) Left SDS-PAGE lanes: Coomassie blue shows HaloTag standard protein samples prepared to test HaloTag ligand-TMR labeling efficiency. “Pre” refers to protein prior to exposure with TMR ligand and bio-beads. “Pre+TMR” refers to the protein after TMR ligand labeling prior to bio-beads. “Post” refers to protein after exposure with TMR ligand and bio-beads. The zoomed-in inset (middle) shows TMR illumination. Dye at the bottom of the gel indicates TMR ligand was added in excess (cyan indicates detector saturation). Rightmost: Integrations of Coomassie blue band intensities were done using ImageJ’s Gel Analysis plug-in. To account for protein that may have nonspecifically bound to the beads, the signal intensities of pre- and post-bead gel bands were compared. (B) The absorbance curve of the post-bead HaloTag standard protein samples. (C) Values used to determine HaloTag ligand-TMR labeling efficiency of the HaloTag standard protein. All Myo10 molecule measurements were adjusted assuming 90% HaloTag ligand-TMR labeling efficiency. (D) In live cells: effect of HaloTag ligand-TMR concentrations on the total intracellular Myo10 signal (left) and total punctate Myo10 signal (right). There was higher, uneven background in the higher HaloTag ligand-TMR live samples that falsely inflated final Myo10 intracellular signal. Plotting only punctate Myo10 signal highlights that signal saturation occurs at 0.5 µM in living cells. (E) In fixed cells: effect of TMR-HaloTag ligand concentrations on the total intracellular Myo10 signal (left) and total punctate Myo10 signal (right). Signal saturation occurs at 0.5 µM in fixed cells.

Added in the revised version: Calculations to convert signal from SDS-PAGE and epifluorescence microscopy to Myo10 molecule counts.

Added in the revised version: Myo10 is irregularly distributed across the plasma membrane.

Added in the revised version: In fixed cells: filopodial Myo10 signal distribution in 150 cells. The center of each Myo10 punctum was plotted after 2-dimensional PCA. Each Myo10 punctum is colored by signal intensity (log-transformed), where magenta = high signal and green = low signal. Cells display periodic stretches of higher Myo10 signal along the cell membrane. Even cells with sparse Myo10 show uneven Myo10 filopodial distribution. Some cells even have membrane patches of no Myo10 signal.

Added in the revised version: Example cell displaying Myo10 density pattern.

Added in the revised version: (A) Cell 14 from Figure 2—figure supplement 1, though raw signal is plotted. 8 values >100,000 raw signal counts not shown. (B) The fixed cell image of (A). Myo10 often concentrates in zones at opposing sides of the cell. Filopodia with high Myo10 signal are next to filopodia also high in Myo10. Actin is labeled with phalloidin-AF633 (magenta). Myo10 is labeled with HaloTag ligand-TMR (green). Scale bar = 10 µm.

Added in the revised version: HaloTag ligand-TMR labeling to determine Myo10 distributions along filopodia.

Added in the revised version: (A) In fixed cells: dim Myo10 puncta in filopodia, identified as those containing <35 Myo10 molecules, were examined. 135 total dim puncta across 26 cells were identified. Of the 135 puncta, 7 puncta were hard to visually interpret, 7 puncta were image noise, and 56 puncta were over-segmented (e.g. not actual distinct puncta but rather part of continuous Myo10 signal in filopodia). The distribution of the remaining 66 true dim puncta is displayed. Min = 6 (95% CI: 5-7), median = 19 (95% CI: 16-22), max = 34 (95% CI: 28-39). Bins = 30. (B) Filopodial localization of dim Myo10 puncta containing <35 molecules. The 66 true dim puncta from (A) were analyzed. (C) Example image demonstrating how Myo10 punctum position was determined in filopodia analyzed in (B). Actin was labeled with phalloidin-AF633 (magenta). HaloTag was labeled with HaloTag ligand-TMR (green). Zoomed-in inset features a filopodium displaying segmented Myo10 puncta localized in the middle (blue arrows) vs. at the tip (orange arrow); zoomed-in inset was contrast adjusted to maximally highlight dim puncta. Scale bar = 5 µm. (D) Log-transformed distribution of the number of Myo10 molecules per filopodium. 8,733 total filopodia across 150 cell images analyzed. Bins = 100. (E) Quantile-quantile (QQ) plot for the log-transformed distribution of Myo10 molecules in filopodia.

Added in the revised version: 96 µM of actin monomers are accessible to Myo10 in a filopodium.

Added in the revised version: Plasma membrane at the filopodial tip can accommodate a portion of Myo10 unbound to actin.

Added in the revised version: Myo10 is found along the filopodium shaft and can undergo a second elongation.

Added in the revised version: (A) Frame stills from Figure 4—video 2. Red arrow points to a Myo10 punctum (red circle) right after emergence from an existing filopodium. Elapsed seconds from the first frame is depicted. Scale bar = 5 µm.

Added in the revised version: Supporting information